Modulating task outcome value to mitigate real-world procrastination via noninvasive brain stimulation

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

The authors report the results of a tDCS brain stimulation study (verum vs sham stimulation of left DLPFC; between-subjects) in 46 participants, using an intense stimulation protocol over 2 weeks, combined with an experience-sampling approach, plus follow-up measures after 6 months.

Strengths:

The authors are studying a relevant and interesting research question using an intriguing design, following participants quite intensely over time and even at a follow-up time point. The use of an experience-sampling approach is another strength of the work.

Comments on revisions:

Overall, I think the authors made many improvements to their manuscript. There are, however, still a number of concerns that first need to be addressed, since it is still not currently possible to fully evaluate the analyses, results, and conclusions presented in the paper. I list these points below:

(1) The authors still use causal language where they must not use causal language. This is true for many places in the manuscript; I am highlighting here just a few places, but the authors nevertheless have to go carefully through the whole manuscript to change these instances.

We sincerely thank the reviewer for this critical and well-taken point. We fully agree that our design (manipulating DLPFC excitability while measuring procrastination, task value, and aversiveness) does not directly measure or manipulate self-control, nor does it rule out alternative neurocognitive mechanisms. Accordingly, we have conducted a comprehensive, line-by-line revision of the entire manuscript to systematically replace causal claims with cautious, hypothesis-consistent language. In response, we have replaced all the wordings that may imply causal inferences, such as impair, cause, boost, by association-consistent phrasing. Furthermore, as you clearly raised below, we explicitly reframed self-control as a hypothesized theoretical construct rather than an empirically verified mediator throughout the whole re-revised manuscript. Please see specific revisions below:

Abstract Section (Page 2, Line 53-59)

“... a mediation analysis indicated a disassociable mechanism: the increase in task outcome value (but not task aversiveness) showed a statistical pattern consistent with accounting for the observed behavioral improvement. In conclusion, these findings are consistent with the hypothesis that enhancing DLPFC function may reduce procrastination by selectively amplifying the valuation of future rewards, not by simply reducing negative feelings about the task.”

Introduction Section (Page 3, Line 83-86)

“... Even worse, chronic procrastination has been consistently associated with poor general health conditions, such as immune system disruption, gastrointestinal disturbance, hypertension and cardiovascular disease (Sirois, 2015; Sirois, 2016).”

Introduction Section (Page 4, Line 143-145)

“Consistent with this framework, the left dorsolateral prefrontal cortex (DLPFC)—a region frequently implicated in value-based decision-making and top-down regulation—has been associated with procrastination. ...”

Introduction Section (Page 4, Line 135-139)

“... Also, given the integrative nature of prefrontal regulatory functions, we hypothesize a third pathway whereby both decreased task aversiveness and increased task-outcome value may jointly contribute to reduced procrastination, potentially reflecting coordinated downstream effects on valuation and affective processing.”

Introduction Section (Page 5, Line 183-186)

“... Thus, this study aims to clarify the brain-behavior association between DLPFC neuromodulation and procrastination, and to test whether observed changes in task valuation and aversiveness are consistent with theoretical models of top-down regulation.”

Results Section (Page 11, Line 534-536)

“... Thus, these findings are consistent with the view that neuromodulation of the left DLPFC is associated with reduced task aversiveness and increased task-outcome value.”

Results Section (Page 11, Line 579-584)

“... In summary, these findings identified a statistical pathway consistent with the theoretical model: neuromodulation of the left DLPFC was associated with increased task-outcome value, which in turn was associated with reduced procrastination.”

Discussion Section (Page 12, Line 617-620)

“On balance, our findings provide evidence consistent with the hypothesis that neuromodulation of the left DLPFC is associated with reduced procrastination, primarily through increasing task-outcome value rather than merely reducing task aversiveness. ...”

Discussion Section (Page 13, Line 664-669)

“... Building on this foundation, among several theoretical interpretations and cognitive pathways, our study showed the one plausible neurocognitive mechanism of procrastination: the cortical excitability of the DLPFC produced by active neuromodulation may engage prefrontal regulatory networks to increase task outcome value, which in turn is associated with reduced procrastination behavior, statistically supporting the theoretical accounts of temporal decision model (TDM, Zhang et al., 2019).”

Discussion Section (Page 14, Line 755-762)

“... Moreover, this study did not collect data for assessing participants' self-control at either baseline or post-neuromodulation. Accordingly, we explicitly note that self-control was not directly measured or manipulated in this study; the observed associations between DLPFC neuromodulation, task-outcome value, and procrastination are consistent with theoretical models positing a role for top-down regulatory processes, but do not constitute direct evidence that self-control mechanisms were engaged. This limitation precludes definitive conclusions about the unique contribution of self-control-related pathways versus alternative neurocognitive mechanisms.”

Some examples:

(a) In response to my comment (1) in the previous round, where the authors adjusted their text, the authors still use causal language in their last sentence "... procrastination behavior has been observed to impair general health..." Unless the cited study truly allowed causal conclusions, the causal language should be removed here as well.

Thank you for pointing out this inappropriate phrasing. As you kindly suggested, we have reworded it as “Even worse, chronic procrastination has been consistently associated with poor general health conditions, such as immune system disruption, gastrointestinal disturbance, hypertension and cardiovascular disease” (Introduction Section, Page 3, Line 83-86).

(b) The authors still make (causal) claims about the involvement of self-control in their observed results. To reiterate from the previous round of revisions: The authors cannot make any strong claims about the role of self-control processes because they do not directly measure self-control nor do they directly manipulate self-control or have a design that would rule out alternative mechanisms other than self-control. Therefore, their claims about self-control have to be toned down. It is laudable that the authors have added a statement towards the end of their discussion about not being able to make strong conclusions about the role of self-control. But the authors need to use similar careful wording not just at the end of the discussion but throughout the manuscript.

We appreciate you reiterating this concern. In the re-revised manuscript, we have thoroughly removed or rewritten all the statements implying causal inferences, and have substantially toned-down claims for the roles of self-control in procrastination reduction from this neuromodulation. Please see instances for what we have replied to the Comment #1.

(i) In the abstract, the authors use the formulation "...conceptualized roles of self-control on procrastination..." -- this wording is still too strong, suggesting that you actually studied self-control.

Thank you for providing this specific instance. This inappropriate sentence has been removed.

(ii) In the introduction (page 4, lines162-169), the way the authors formulate these sentences suggests that they directly measured self-control. Again, the authors need to make it explicit that they are not directly measuring self-control but its hypothesized down-stream consequences on valuations/behavior.

Many thanks. This statement has been removed, and we have reworded it as “... Thus, this study aims to clarify the brain-behavior association between DLPFC neuromodulation and procrastination, and to test whether observed changes in task valuation and aversiveness are consistent with theoretical models of top-down regulation.” (Introduction Section, Page 5, Line 183-186).

(iii) In the discussion, for example, on page 11, lines 555 and following, the authors write: "One major contribution this study has made is to disentangle the neurocognitive mechanism of procrastination by demonstrating that self-control could increase task-outcome value so as to reduce procrastination."

As you kindly instructed, we have rewritten this statement as “One contribution of this study is to provide empirical evidence partially consistent with the temporal decision model (TDM), showing that increased task-outcome value—rather than decreased task aversiveness—was statistically associated with reduced procrastination following DLPFC neuromodulation.” (Introduction Section, Page 12, Line 625-628), which no longer implies any conclusions for the role of self-control in this study.

Again, please be aware that you are NOT demonstrating that self-control does anything, since you only measure procrastination rates, outcome values, and task aversiveness. It is possible that mechanisms other than self-control might be relevant for this. Perhaps neuromodulation directly increases outcome values, without involvement of self-control processes. You simply cannot know that and therefore you cannot make those claims in the form that you are making them. You can write that the observed results are consistent with the idea that neuromodulation might have had an effect on self-control and this in turn might have affected outcome values. But you also need to make it explicit that, to substantiate these claims, you would need more direct evidence that indeed self-control was involved. These more careful formulations would not at all reduce the value of your work, but indeed they would rather demonstrate your carefulness in interpreting the results you obtained.

We sincerely thank the reviewer for this exceptionally clear and constructive guidance. We fully agree that our study design does not measure or manipulate self-control, and therefore we cannot demonstrate that self-control processes are causally involved in the observed effects. As you correctly note, it is entirely possible that neuromodulation directly modulates outcome valuation or engages alternative neurocognitive pathways (e.g., attentional allocation, feedback learning, or affective processing) without invoking self-control mechanisms.

In direct response, as we replied above, we have completely rewritten the whole revised manuscript to remove any assertions that we “identified” a role of self-control. The revised text now explicitly states as follow: (1) our findings merely are consistent with the theoretical hypothesis that DLPFC neuromodulation might engage prefrontal self-regulatory functions, which in turn influence outcome valuation; (2) we explicitly acknowledge that substantiating this specific pathway would require more direct evidence. Rather than single sentence, we have applied this careful, hypothesis-consistent framing systematically across the Abstract, Introduction, Results, and Discussion. As you suggested, these revisions more accurately reflect the interpretative boundaries of our data and demonstrate our commitment to rigorous, transparent scientific reporting. Please see specific cases for this revision above.

(2) I am still puzzled by the power analysis. In the text, you write that a sample size of 18 participants (i.e., 9 per group) would be sufficient to achieve 80% power. I still feel this seems far too optimistic and hard to believe, but that is not my point here. While in the text, you write that you need 18 participants, the G*power output seems to suggest a sample size of 34, not 18. Why this contradiction? Or is it not contradictory? If it is not, then please explain it more fully.

We appreciate you pointing out this critical typo. In the last round of revision, we mean that 18 participants per group are required to achieve at least 80% statistical power, rather than a total sample size, as shown by the GPower software. We are sorry for this critical typo to confuse you. As you correctly pointed out, the GPower indicated that the minimum sample size to reach 80% power is 34 (i.e., 17 per group). Thus, we selected 36 (i.e., 18 per group) participants as minimum sample size in case of potential drop-out. We have thoroughly corrected this typo, and double-checked no such numeric issues:

Methods Section (Page 5, Line 234-237)

“... statistical power was predetermined by G*Power at a relatively medium effect size (1-β err prob = 0.80, f = 0.25), indicating the total sample size at 34 (17 per group) to reach acceptable power. To account for potential attrition, we determined to recruit 36 participants, at least.”.

(3) I have several comments about the mixed-effects analysis.

First of all, I want to thank the authors for adding more details, things have become much clearer now. However, I still have a few questions and comments related to these analyses:

(a) The variable Emotions was within-subjects, as far as I understood. Accordingly, Emotions should most likely be modelled with random slopes varying over participants (in addition to being modelled as a fixed effect).

We thank you raising this reasonable concern on the mixed-effect linear modeling. Yes, the Emotions reflect daily baseline affect, which is modeled as covariates of no interests to adjust for daily emotional fluctuation (if any). In this vein, this baseline emotion score is included for each participant across all the sessions. Therefore, it should be modeled with random slopes as you assumed indeed.

In response, we have remodeled this mixed-effects analysis by including the daily baseline emotion as random slopes varying over participants. After centering the variables, we estimated the revised models as “Procrastination Rate ~ Group * Treatment day + Age + Gender + SES + Emotions + (1 + Treatment day + Emotions || SubjectID)” and “Task execution willingness ~ Group * Treatment day + Age + Gender + SES + Emotions + (1 + Treatment day + Emotions || SubjectID)”. Notably, as you correctly assumed, fitting this complicated random-effect structure is likely to result in convergence failure, given the limited sample size in the present study. Therefore, we hypothesized the independence among random effects for model simplification. Consistent with this assumption, model comparisons indicated that the simplified models fit better than original ones (Procrastination Rate model, ∆AIC = -1.0, ∆BIC = -11.8, LRT, χ2(3) = 5.04, p = .17; Task-execution willingness model, ∆AIC = -5.5, ∆BIC = -16.4, LRT, χ2(3) = 0.51, p = .91).

Taken together, as you kindly suggested, we have rebuilt the mixed-effect models by adding daily baseline emotion as random slopes varying over participants, and have demonstrated the consistent findings with the original one:

Methods Section (Page 8-9, Line 412-419)

“... Given the risks of convergence failure with the two correlated random-effects structure (i.e., treatment days and self-reported emotions), we hypothesized that the random effects are independent, leading to model simplification. Consistent with this assumption, model comparisons favored the simplified independent structure over the full correlated structure for both outcomes. For the actual procrastination model, the simplified model showed lower AIC (∆ = -1.0) and BIC (∆ = -11.8), with a non-significant likelihood ratio test (χ2 (3) = 5.04, p = .17). For the task-execution willingness model, the simplified model was also preferred (∆AIC = -5.5, ∆BIC = -16.4; LRT: χ2 (3) = 0.51, p = .91).”

Results Section (Page 9-10, Line 469-489)

“For procrastination willingness, results showed a statistically significant interaction effect between multi-session neuromodulations and groups (β = -7.84, SE = 1.80, t = -4.36, DF = 45.6, p < .001; Fig. 3A). In the post-hoc simple effect analysis, it demonstrated a significantly increased task-execution willingness (i.e., decreased procrastination willingness) after neuromodulation in the active neuromodulation group (NM-before: 35.65 ± 30.21, NM-after: 80.43 ± 19.92, Mean Diff = 41.79, SE = 7.58, DF = 103.4, t.ratio = 5.51, p < .0001, Tukey correction), but no such effects were identified in the sham control group (SC-before: 37.57 ± 26.46, SC-after: 47.35 ± 30.49, Mean Diff = 2.58, SE = 7.56, DF = 96.8, t.ratio = 0.34, p = .73, Tukey correction) (Fig. 3B-C). A linear uptrend for task-execution willingness was further observed across multiple sessions in the active NM group, indicating gradually increasing neuromodulation effects (Fig. 3D; p < .01, Mann-Kendall test). For actual procrastination behavior, changes to actual procrastination rates across all the sessions have been detailed in the Fig. 3E. Similarly, a statistically significant interaction effect was identified here (β = -7.37, SE = 2.40, t = -3.02, DF = 46.6, p = .004), and the simple effect analysis further revealed decreased actual procrastination rates after ms-tDCS in the active neuromodulation group (NM-before: 56.74 ± 39.10, NM-after: 0.00 ± 0.00, Mean Diff = 44.40, SE = 9.36, DF = 110.0, t.ratio = 4.74, p < .0001, Tukey correction), but no such prominent changes found in the sham control group (SC-before: 46.47 ± 40.76, SC-after: 33.35 ± 37.82, Mean Diff = 7.53, SE = 9.28, DF = 102.0, t.ratio = 0.81, p = .42, Tukey correction) (Fig. 3F-G).”

(b) The analyses still cannot fully be evaluated as I cannot access the scripts and data. The authors mention that the scripts and data should be available via a link they provide (https://doi.org/10.57760/sciencedb.35140). However, when I try to access these materials via this link, no page opens; it seems the link is dead?

Thank you very much for bringing this case to us. We checked this link and found it to be still active.

To ensure accessibility for your evaluation, we have uploaded scripts and data into this online submission system. Please do let us know if you are still unable to access them. We are glad to send them to you by other available pathways. This link is a private access to you, and the repository would be openly available for other users upon the final publication.

(c) What are the results and conclusions if you do not include the covariates of no interest? I.e., please re-run your main models without age, gender, SES, Emotions.

Thank you for raising this question. As you clearly instructed, we have rerun main models without all those covariates. As shown in the table below, the results for the key predictors of interest (Group, Treatment day, and their interaction) remained largely unchanged in terms of effect size, direction, and statistical significance:

(d) The authors mention that they use GLMMs, which would suggest generalized mixed-effects models, but they do not describe what family/distribution they used. Since they mention lmerTest and seem to report F-tests, my guess is that they used Gaussian models. However, both their DVs (procrastination rates and their ratings) are bounded variables and at least procrastination rates hit the lower boundary. That can mean that their analyses suffer from inflated Type 1 and/or Type 2 rates. Therefore, please repeat the analyses with an appropriate generalized mixed-effects model (perhaps a beta regression type of model?).

We are very grateful to you for raising this crucial statistical point. As you correctly pointed out, we used the Gaussian distribution in estimating this model. We are sorry to confuse you due to the absence of reporting family/distribution we used. In the original manuscript, we meant “general” linear mixed-effect model, rather than “generalized” one. As you clearly and correctly raised, procrastination rates and willingness are technically bounded, and that procrastination rates frequently reached the lower boundary (0%) in the present study, which are in high risks to be inflated for Type 1 and/or Type 2 error.

Thus, as you kindly suggested, a beta family distribution with logit function is used to reanalyze those main effects of interest. Results are tabulated in Author response table 2.

These convergent results confirm that the critical main effect (i.e., Group and Treatment day) and their interaction remain statistically significant across distributional specifications, and that our primary conclusions are not artifacts of the Gaussian assumption. Taken them together, as you kindly suggested, we have repeated the analyses with beta regression family distribution, which replicated our main findings, potentially supporting their statistical robustness.

Following your suggestion, we have added those results derived from such sensitivity analyses into the revised manuscript:

Methods Section (Page 9, Line 430--436)

“... To examine whether our findings were sensitive to the distributional assumptions of the dependent variables, we re-analyzed the main models using an alternative distributional specification. Given that both procrastination rates (ranging from 0% to 100%) and task-execution willingness (measured on a 0-100 visual analog scale) are bounded continuous outcomes, and that procrastination rates frequently reached the lower boundary (0%) in the present study, a Beta regression model with a logit link function was employed for a sensitivity analysis.”

Results Section (Page 10, Line 506-512)

“... Furthermore, as a sensitivity analysis, we reran the main LMMs using Beta regression distribution with a logit link function, which is appropriate for the both bounded outcomes mentioned above (i.e., procrastination rate and procrastination willingness). The main effects (i.e., Group and Treatment day) and their interaction remained significant for both procrastination rate and willingness (see SI Results and Tab. S5), confirming that our findings are robust to alternative distributional assumptions.”

(e) When reporting the results of the mixed-effects models, the authors report the regression coefficient, standard error, DFs and p value, but not the actual test statistic. Please add the information about the test statistic and report all degrees of freedom (in case of F tests that would be the degrees of freedom of the test and the residual degrees of freedom).

We truly thank you for this nuanced reminder. As you suggested, we have added actual test statistics, including t-values and all degrees of freedom (DF). Please see specific instances below:

Results Section (Page 9-10, Line 469-489)

“For procrastination willingness, results showed a statistically significant interaction effect between multi-session neuromodulations and groups (β = -7.84, SE = 1.80, t = -4.36, DF = 45.6, p < .001; Fig. 3A). In the post-hoc simple effect analysis, it demonstrated a significantly increased task-execution willingness (i.e., decreased procrastination willingness) after neuromodulation in the active neuromodulation group (NM-before: 35.65 ± 30.21, NM-after: 80.43 ± 19.92, Mean Diff = 41.79, SE = 7.58, DF = 103.4, t.ratio = 5.51, p < .0001, Tukey correction), but no such effects were identified in the sham control group (SC-before: 37.57 ± 26.46, SC-after: 47.35 ± 30.49, Mean Diff = 2.58, SE = 7.56, DF = 96.8, t.ratio = 0.34, p = .73, Tukey correction) (Fig. 3B-C). A linear uptrend for task-execution willingness was further observed across multiple sessions in the active NM group, indicating gradually increasing neuromodulation effects (Fig. 3D; p < .01, Mann-Kendall test). For actual procrastination behavior, changes to actual procrastination rates across all the sessions have been detailed in the Fig. 3E. Similarly, a statistically significant interaction effect was identified here (β = -7.37, SE = 2.40, t = -3.02, DF = 46.6, p = .004), and the simple effect analysis further revealed decreased actual procrastination rates after ms-tDCS in the active neuromodulation group (NM-before: 56.74 ± 39.10, NM-after: 0.00 ± 0.00, Mean Diff = 44.40, SE = 9.36, DF = 110.0, t.ratio = 4.74, p < .0001, Tukey correction), but no such prominent changes found in the sham control group (SC-before: 46.47 ± 40.76, SC-after: 33.35 ± 37.82, Mean Diff = 7.53, SE = 9.28, DF = 102.0, t.ratio = 0.81, p = .42, Tukey correction) (Fig. 3F-G).”

(f) Thank you for adding the analysis where you remove the last two sessions. But currently you present them in the manuscript without explaining/motivating why you do this. Please add this motivation, as otherwise it will be puzzling for the reader why you conduct these analyses.

Thank you for this very practical requirement to clarify the motivation of reanalyzing main models from removing the last two sessions. Please see specific explanation as follow:

Results Section (Page, Line 499-506)

“... To systematically test whether such effects are biased by extreme data points or patterns, we reran the main LMMs by iteratively removing data from the last two sessions, which showed extraordinarily high effectiveness from neuromodulation (e.g., all the participants in the neuromodulation group had no actual procrastination behavior in session #6 and #7). Results showed the significant group*neuromodulation sessions interaction effects across all those nested models (removing session #6, #7 or both, all p < .05; see SI Results and Tab. S3-4), potentially indicating a statistical robustness from the data pattern.”

(4) Mediation analysis

In your manuscript, you present some mediation analyses. Please be aware that such mediation analyses cannot establish causality and they suffer from extremely high Type 1 error rates (see, e.g., https://datacolada.org/103). My suggestion would be to completely remove all mediation analyses. However, if you want to keep them, then you need to be extremely careful in how you present the results. You need to explicitly mention that you cannot derive any causal conclusions from them and that simulation studies have shown that such mediation analyses suffer from extremely high Type 1 errors.

We sincerely thank you for this exceptionally important methodological guidance. We fully agree that mediation analyses, especially those based on observational measures rather than experimentally manipulated mediators, cannot establish causal pathways and are susceptible to inflated Type 1 error rates, as rigorously demonstrated in recent simulation studies (https://datacolada.org/103).

As you kindly suggested, please allow us to retain those mediation analyses upon explicitly highlighting that this mediation statistical model cannot generate any causal conclusions. In response, we have systematically replaced all instances of “causal mediation” with “statistical mediation” or “exploratory mediation analysis”, and removed causal verbs (e.g., “depends on”, "drives", "explains") in favor of association-consistent phrasing (e.g., “is statistically mediated”, “aligns with the hypothesis that”) throughout the abstract, introduction, methods, results and discussion sections. Furthermore, in the Discussion section, we explicitly reiterated the limitations of extending this mediation associations to causal conclusions. Please see specific modifications underneath:

Abstract Section (Page 2, Line 52-56)

“... While the intervention is significantly associated with both decreased task aversiveness and increased perceived task outcome value, a mediation analysis indicated a disassociable mechanism: the increase in task outcome value (but not task aversiveness) showed a statistical pattern consistent with accounting for the observed behavioral improvement.”

Methods Section (Page 9, Line 448-453)

“... As these mediation analyses are based on observational measures rather than experimentally manipulated mediators, they do not establish causal pathways. Simulation studies have shown that such analyses can suffer from inflated Type 1 error rates. Results should therefore be interpreted as hypothesis-generating and statistically consistent with the proposed theoretical model, rather than as confirmatory evidence of causal mechanisms.”

Methods Section (Page 9, Line 438-440)

“... the Quasi-Bayesian mediation analysis was used to model the association between the effects of tDCS, task aversiveness/outcome and decreased procrastination.”

Results Section (Page 11, Line 568-573)

“As an exploratory analysis, results indicated that increased task outcome value was associated with changes in the task-execution willingness (δ = 21.73, p < .01; ζ = 11.25, p = .07, ρ = 32.99, p < .01, simulation = 1,000; see Fig. 5A) and real-world procrastination (δ = 30.75, p < .01; ζ = 3.05, p = .52, ρ = 33.81, p < .01, simulation = 1,000; see Fig. 5B), in the context of ms-tDCS neuromodulation. ...”

Results Section (Page 11, Line 582-584)

“... Nevertheless, all mediation findings are now explicitly labeled as “exploratory” and framed as quantifying statistical associations consistent with the TDM pathway, not causal mediation.”

Discussion Section (Page 14, Line 747-755)

“Notably, we explicitly acknowledge that exploratory Quasi-Bayesian mediation analyses, based on observational measures rather than experimentally manipulated mediators, cannot establish causal pathways and are susceptible to inflated Type 1 error rates as demonstrated in recent simulation studies. These findings should be interpreted strictly as hypothesis-generating and statistically consistent with the proposed theoretical model, rather than as confirmatory evidence of causal mechanisms. Substantiating the precise neurocognitive pathway will require future studies employing stronger causal designs, such as experimental manipulation of task valuation or longitudinal cross-lagged modeling. ...”

As an example (but the mediation results are mentioned in several places, for example, also in the abstract): On page 10, lines 501-503: What you can causally conclude is that neuromodulation affects your measured variables (outcome values, procrastination rates, task aversiveness), but you cannot conclude that the effect of neuromodulation on procrastination rates causally operates via outcome values. Thus, please adjust the formulation accordingly. The same applies to the mediation section that follows right afterwards (page 10, lines 505-522).

Thank you for offering those specific instances. As we replied above, they have been revised accordingly:

Results Section (Page 11, Line 557-559)

“... Collectively, these findings provide statistical evidence consistent with the hypothesis that the outcome-value pathway may contribute to procrastination reduction.”

Results Section (Page 11, Line 563-582)

“Increased task outcome value is specifically associated with reduced procrastination in the context of neuromodulation

To explore the potential neurocognitive pathways of procrastination, the Quasi-Bayesian mediation analysis was undertaken, with increased task outcome value as a statistically mediated variable. As an exploratory analysis, results indicated that increased task outcome value was associated with changes in the task-execution willingness (δ = 21.73, p < .01; ζ = 11.25, p = .07, ρ = 32.99, p < .01, simulation = 1,000; see Fig. 5A) and real-world procrastination (δ = 30.75, p < .01; ζ = 3.05, p = .52, ρ = 33.81, p < .01, simulation = 1,000; see Fig. 5B), in the context of ms-tDCS neuromodulation. To ensure the statistical robustness and specificity of these findings, the sensitivity analysis was implemented by changing sampling parameters and outcome variables. By doing so, those findings were validated statistically robust, as shown by replicated observations across bootstrapping sampling subsets (see SI Results and Tab. S6-7). Moreover, the results of the control analysis further validated the specificity of these findings by showing a null statistically mediated effect of this model to predict one’s task aversiveness (see SI Results and Tab. S8). In summary, these findings identified a statistical pathway consistent with the theoretical model: neuromodulation of the left DLPFC was associated with increased task-outcome value, which in turn was associated with reduced procrastination.”

(5) In the introduction, the authors introduce several theoretical procrastination frameworks (TMT, mood repair, TDM). Do the results of the current paper help to decide which framework might be the most appropriate, at least for the authors data set? It might be of interest to address this explicitly.

We do thank the reviewer for this insightful theoretical question. We agree that explicitly positioning our findings within the broader theoretical landscape strengthens the conceptual contribution of our work. Upon careful consideration, we believe that our results provide the strongest empirical support for the TDM over alternative frameworks (TMT, mood repair). TDM uniquely posits procrastination as contingent on the dynamic trade-off between task aversiveness and task-outcome value. In the present study, neuromodulation was identified to be associated with both pathways but only increased outcome value statistically predicted reduced procrastination. This aligns precisely with TDM’s hypothesis that value-based processes may dominate aversiveness-avoidance processes in driving behavioral change. Neither TMT (which emphasizes temporal discounting of utility per se) nor the mood repair perspective (which prioritizes short-term affect regulation) explicitly predicts this dissociable pattern. As you kindly suggested, we have extended the Discussion section to explicitly contextualize our findings into this theoretical landscape (Discussion Section, Page 13, Line 669-687).

(6) The language is sometimes hard to understand and seems in quite some places grammatically incorrect. Thus, I think the paper would profit very much from thorough English proofreading.

We sincerely thank you for this practical suggestion. We fully agree that the original manuscript contained grammatical inaccuracies and awkward phrasing that could hinder readability. In response, we have engaged a professional academic editing service to thoroughly proofread and polish the entire manuscript. All sentences have been revised for grammatical correctness, syntactic clarity, and academic tone, while strictly preserving the original scientific meaning and technical terminology. We believe this language improvements have substantially enhanced the readability and overall quality of the paper.

Reviewer #2 (Public review):

Summary:

Chen and colleagues conducted a cross-sectional longitudinal study, administering high-definition transcranial direct stimulation (HD-tDCS) targeting the left DLPFC to examine the effect of HD-tDCS on real-world procrastination behavior. They find that seven sessions of active neuromodulation to the left DLPFC elicited greater modulation of procrastination measures (e.g., task-execution willingness, procrastination rates, task aversiveness, outcome value) relative to sham. They show that HD-tDCS reduces task aversiveness and increases task-execution willingness on real-world tasks as quantified by intensive experience sampling methods, providing causal evidence for the role of DLPFC in modulating contextual features to delaying or completing one's goals.

Strengths:

This is a well-designed protocol with rigorous administration of high-definition transcranial direct current stimulation across multiple sessions. The intensive experience sampling approach which probes and assesses self-relevant task goals is innovative and aims to address an important question regarding the specific role of DLPFC in modulating specific features of chronic procrastination behavior (e.g., task-execution willingness, task aversiveness).

The quantification of task aversiveness through AUC metrics is a clever approach to account for the temporal dynamics of task aversiveness, which is notoriously difficult to quantify.

Weaknesses:

While the findings that neurostimulation reduces procrastination behavior is compelling, there remain several alternative interpretations for these effects. For example, it could be that the task-execution willingness isn't increased per se, but rather that the goal completion becomes more valuable as participants learn from feedback or become more aware of their successful attainment of or failure to complete task goals. It is unclear whether the effects could be driven by improved working memory or attention to the reported tasks (and this limitation is addressed by the authors). In short, it is also difficult to examine the temporal dynamics of how these goals are selected across time.

We sincerely thank you for raising these thoughtful and methodologically important points. We fully agree that the observed reductions in procrastination could reflect multiple neurocognitive pathways beyond the value-based mechanism emphasized in our primary analysis.

In response, we have thoroughly removed claims on the “unique mechanism” of value-based pathways to procrastination reduction, and fully substituted language implying exclusive mediation by “value amplification” with more cautious phrasing (e.g., “statistically consistent with a value-based pathway”; “one plausible mechanism among several processes”). In the revised manuscript, we reiterated that the pattern of results, showing increased outcome value predicting reduced procrastination while decreased aversiveness did not, aligned with the Temporal Decision Model, yet does not rule out concurrent contributions from attention, learning, or executive processes. To clearly bring this interpretative boundary of our primary findings for audiences, we have explicitly warranted such cautions in the Discussion Section. Please see specific modifications underneath:

Discussion Section (Page 13, Line 664-669)

“... Building on this foundation, among several theoretical interpretations and cognitive pathways, our study showed the one plausible neurocognitive mechanism of procrastination: the cortical excitability of the DLPFC produced by active neuromodulation may engage prefrontal regulatory networks to increase task outcome value, which in turn is associated with reduced procrastination behavior, statistically supporting the theoretical accounts of temporal decision model (TDM, Zhang et al., 2019). ”

Discussion Section (Page 13, Line 676-687)

“... Despite statistically supporting the TDM, we acknowledge that alternative neurocognitive mechanisms could contribute to the observed reductions in procrastination. For instance, repeated exposure to the experience-sampling protocol may have enhanced participants’ awareness of task progress or facilitated feedback-based learning, thereby increasing the subjective value of goal completion independent of DLPFC neuromodulation. Similarly, improvements in working memory for task maintenance, attentional allocation to reported goals, or strategic shifts in goal selection across sessions could plausibly mediate the intervention effects. While our double-blind, sham-controlled design and inclusion of daily emotional covariates help mitigate some non-specific confounds, the present study did not incorporate direct measures of these alternative processes. Consequently, we cannot definitively isolate the value-based pathway posited by the TDM from concurrent contributions of attention, learning, or executive functions.”

It is unclear whether the current evidence support long-retention of this neurostimulation intervention. The study includes one 6-month timepoint after the study to examine the long-term retention of the neural stimulation effect. Future studies that evaluate the long-term effects across multiple time points would strengthen the evidence for the robustness of this intervention.

We genuinely appreciate you for this insightful and methodologically reasonable point. We fully agree that a single 6-month follow-up assessment, while valuable, provides only preliminary evidence for long-term retention, and that multiple follow-up timepoints would substantially strengthen claims about the durability of neuromodulation effects. To carefully address this point, we have rephrased the “long-term retention” as “long-term after-effects” throughout the whole revised manuscript, and overall toned down the claims on the retention effects. Moreover, this limitation has been explicitly elaborated in the Discussion section:

Abstract Section (Page 2, Line 49-50)

“... we assessed the effect of anodal HD-tDCS on real-world procrastination behavior at offline after-effect (2-day interval) and long-term after-effect (6-month follow-up).”

Results Section (Page 12, Line 606-608)

“... Therefore, beyond short-term effects, the benefits of ms-tDCS neuromodulation on reducing procrastination were still detectable at a 6-month follow-up, providing preliminary evidence consistent with long-term after-effects.”

Discussion Section (Page 14, Line 721-725)

“... Thus, the detectable effects at 6 months are consistent with the hypothesis that repeated neuromodulation may induce neuroplastic changes in the DLPFC that support sustained behavioral change. However, we explicitly note that a single follow-up timepoint cannot establish the stability or trajectory of these effects; future studies with multiple longitudinal assessments are required to substantiate claims about long-term retention.”

Discussion Section (Page 15, Line 771-775)

“... Finally, while our 6-month follow-up provides preliminary evidence for sustained effects, the use of a single follow-up timepoint limits our ability to characterize the temporal trajectory of retention. Future studies incorporating multiple follow-up assessments (e.g., 1-month, 3-month, 6-month, 12-month) would strengthen evidence for the robustness and durability of this intervention.”

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Please see my detailed comments above (6 points; several of them with subpoints a, b, c, etc).

Thank you for listing those specific and helpful recommendations above. Please see our detailed response posed above, point-by-point.

Changes from the earlier version to the revised version

Abstract

Earlier version: Procrastination represents one of the most prevalent behavioral problems affecting individual health and societal productivity. Although it is often conceptualized as a form of self-control failure, its underlying neurocognitive mechanisms are poorly understood. A leading model posits that procrastination arises from imbalanced competing motivations: the avoidance of negative task aversiveness and the pursuit of positive task outcomes, yet this theoretical framework has not fully validated in real-world settings and not applied effectively to guide interventions. Here, we addressed this gap with a double-blind, randomized controlled trial. We applied seven sessions of high-definition transcranial direct current stimulation (HD-tDCS) to the left dorsolateral prefrontal cortex (DLPFC), a key region for self-control, in chronic procrastinators. Using the intensive experience sampling method (iESM), we assessed the effect of anodal HD-tDCS on real-world procrastination behavior at offline after-effect (2-day interval) and long-term retention (6-month follow-up). We found that this neuromodulation produced a lasting reduction in real-world procrastination, with effects sustained at a 6-month follow-up. While the intervention both decreased task aversiveness and increased perceived task outcome value, causal mediation analysis revealed a striking mechanism: the increase in task outcome value (but not task aversiveness) sufficiently mediated the entire behavioral improvement. In conclusion, these findings provide causal evidence that enhancing DLPFC function mitigates procrastination by selectively amplifying the valuation of future rewards, not by simply reducing negative feelings about the task. This establishes a precise, value-driven neurocognitive pathway to account the conceptualized roles of self-control on procrastination, and potentially offers a validated, theory-driven strategy for interventions.

Revised version: Procrastination represents one of the most prevalent behavioral problems associated with individual health and societal productivity. Despite its high prevalence and substantial impact on daily functioning, its underlying neurocognitive mechanisms remain poorly understood. A leading model posits that procrastination arises from imbalanced competing motivations: the avoidance of negative task aversiveness and the pursuit of positive task outcomes, yet this theoretical framework has not been fully validated in real-world settings and not applied effectively to guide interventions. Here, we addressed this gap with a double-blind, randomized controlled trial. We applied seven sessions of high-definition transcranial direct current stimulation (HD-tDCS) to the left dorsolateral prefrontal cortex (DLPFC) in chronic procrastinators. Using the intensive experience sampling method (iESM), we assessed the effect of anodal HD-tDCS on real-world procrastination behavior at offline after-effect (2-day interval) and long-term after-effect (6-month follow-up). We found that this neuromodulation produced a lasting reduction in real-world procrastination, with effects sustained at a 6-month follow-up. While the intervention is significantly associated with both decreased task aversiveness and increased perceived task outcome value, a mediation analysis indicated a disassociable mechanism: the increase in task outcome value (but not task aversiveness) showed a statistical pattern consistent with accounting for the observed behavioral improvement. In conclusion, these findings are consistent with the hypothesis that enhancing DLPFC function may reduce procrastination by selectively amplifying the valuation of future rewards, not by simply reducing negative feelings about the task. These results align with established decision-theoretic frameworks and suggest a targeted, theory-informed avenue for future behavioral interventions.

Introduction

Earlier version: Procrastination is increasingly becoming a prevalent behavioral problem around the world, which reflects the irrational voluntary postponement of scheduled tasks albeit being worse off for such delays (Blake, 2019; Steel, 2007). In epidemiological investigations, more than 15% of adults were identified as having chronic procrastination problems, and the situation for students was worse as 70-80% of undergraduates engaged in procrastination (American College Health Association, 2022; Ferrari et al., 2005). Moreover, the behavioral genetic evidence indicates a certain heritability of procrastination in human beings as well (Gustavson et al., 2017; Gustavson et al., 2014, 2015). In addition to its prevalence, the undesirable associations between procrastination behavior and health also warrant caution. There is cumulative evidence to show the close associations between procrastination behavior and work performance, financial status, interpersonal relationships, and subjective well-being (Ferrari, 1994; Pychyl & Sirois, 2016; Steel et al., 2021). Further, as prospective cohort studies indicate, many mental health problems emerge alongside procrastination, particularly sleep problems, depression, and anxiety (Hairston & Shpitalni, 2016; Johansson et al., 2023). Even worse, chronic procrastination has been observed to impair general health, as manifested by the close associations with immune system disruption, gastrointestinal disturbance, as well as a high risk of hypertension and cardiovascular disease (Sirois, 2015; Sirois, 2016). Thus, given these critical ramifications, considerable efforts have been devoted to delve into why we procrastinate irrationally.

Revised version: Procrastination is increasingly becoming a prevalent behavioral problem around the world, which reflects the irrational voluntary postponement of scheduled tasks albeit being worse off for such delays (Blake, 2019; Steel, 2007). In epidemiological investigations, more than 15% of adults were identified as having chronic procrastination problems, and the situation for students was worse as 70-80% of undergraduates engaged in procrastination (American College Health Association, 2022; Ferrari et al., 2005). Moreover, the behavioral genetic evidence indicates a certain heritability of procrastination in human beings as well (Gustavson et al., 2017; Gustavson et al., 2014, 2015). In addition to its prevalence, the undesirable associations between procrastination behavior and health also warrant caution. There is cumulative evidence to show the close associations between procrastination behavior and work performance, financial status, interpersonal relationships, and subjective well-being (Ferrari, 1994; Pychyl & Sirois, 2016; Steel et al., 2021). Further, as prospective cohort studies indicate, many mental health problems emerge alongside procrastination, particularly sleep problems, depression, and anxiety (Hairston & Shpitalni, 2016; Johansson et al., 2023). Even worse, chronic procrastination has been consistently associated with poor general health conditions, such as immune system disruption, gastrointestinal disturbance, hypertension and cardiovascular disease (Sirois, 2015; Sirois, 2016). Thus, given these critical ramifications, considerable efforts have been devoted to delve into why we procrastinate irrationally.

To probe why we procrastinate irrationally, researchers have built upon theoretical bases of procrastination from different perspectives. For instance, Steel (2007) pioneered a promising temporal motivation theory (TMT) to explicate procrastination as the failure of self-regulation. This theory suggests that individuals would procrastinate as task utility is devalued in the far future (Steel & König, 2006). Based on insights into emotional regulation, the mood repair perspective provides another explanation to elucidate that procrastination is due to the failure of self-regulation to give priority to short-term mood repair caused by doing the task rather than long-term task reward (Sirois & Pychyl, 2013). Recently, the temporal decision model (TDM) of procrastination further provides an integrative framework to explain the procrastination decisions, which highlights that procrastination is contingent on the trade-off between task aversiveness and task-outcome value (Zhang & Feng, 2020; Zhang, Liu, et al., 2019). Task aversiveness reflects how unpleasant individuals perceive tasks to be, with more unpleasant feelings making procrastination more likely (Zhang & Feng, 2020). Task outcome value indicates how much it is worth as we evaluate the benefits it provides us (e.g., keeping body health) once we complete the task before the deadline (e.g., doing scheduled exercise) (Zhang & Feng, 2020). If the task aversiveness is overvalued in this trade-off, the decision to postpone tasks would be made consistently.

Earlier version: It is worthwhile to note that this trade-off and resultant decisions to procrastinate are predominantly influenced by self-control capability, with higher self-control associated with less procrastination behavior (Blake, 2019; Ramzi & Saed, 2019; Zhao et al., 2021). Thus far, there have been three promising pathways attempting to clarify how self-control works in reducing procrastination: one for decreasing task aversiveness, the other one for increasing task-outcome value, and the last one for both (Zhang & Feng, 2020). As identified by both behavioral and neural evidence, procrastinators consistently report high task aversiveness when receiving scheduled tasks, and are more likely to postpone tasks so as to devalue negative task aversiveness (Blunt & Pychyl, 2000; Zhang et al., 2021). Meanwhile, the self-control was observed as powerful for downstream regulation of negative emotional stimuli (Paschke et al., 2016; Tice & Bratslavsky, 2000). Therefore, the first possible pathway is to enhance self-control by facilitating emotional regulation towards task aversiveness (Eckert et al., 2016). On the other hand, as a value-based decision, procrastination behavior is contingent on the evaluation of future outcomes (Rebetez et al., 2016; Zhang, Becker, et al., 2019). Existing evidence has shown that procrastinators generally underestimate the task-outcome value (H. Wu et al., 2016; Zhang et al., 2021). In this vein, it is hard to generate the motivation to take immediate action (Taura et al., 2015). Notably, increasing the value of future rewards has been found effective in making individuals inclined to pursue future outcomes by strengthening self-control (Cho et al., 2015; Kelley et al., 2018). Thus, another pathway worth putting forward is that procrastination behavior could be shaped by increasing future task-outcome value. Also, given the multifaceted roles of self-control, we hypothesize a third pathway whereby both decreased task aversiveness and increased task-outcome value contribute to determining procrastination simultaneously through the exertion of self-control.

Revised version: Theoretical frameworks posit that this trade-off and subsequent procrastination decisions may be modulated by individual differences in top-down regulatory capacity, with higher capacity typically associated with less procrastination (Blake, 2019; Ramzi & Saed, 2019; Zhao et al., 2021). Thus far, theoretical models have proposed three potential pathways through which regulatory processes might contribute to reduced procrastination: one for decreasing task aversiveness, the other one for increasing task-outcome value, and the last one for both (Zhang & Feng, 2020). As identified by both behavioral and neural evidence, procrastinators consistently report high task aversiveness when receiving scheduled tasks, and are more likely to postpone tasks so as to devalue negative task aversiveness (Blunt & Pychyl, 2000; Zhang et al., 2021). Meanwhile, priori literature suggests that top-down regulatory processes may facilitate the downstream modulation of negative emotional stimuli (Paschke et al., 2016; Tice & Bratslavsky, 2000). Therefore, one hypothesized pathway involves the downstream modulation of task aversiveness (Eckert et al., 2016). On the other hand, as a value-based decision, procrastination behavior is contingent on the evaluation of future outcomes (Rebetez et al., 2016; Zhang, Becker, et al., 2019). Existing evidence has shown that procrastinators generally underestimate the task-outcome value (H. Wu et al., 2016; Zhang et al., 2021). In this vein, it is hard to generate the motivation to take immediate action (Taura et al., 2015). Notably, increasing the value of future rewards has been found effective in making individuals inclined to pursue future outcomes by potentially engaging prefrontal regulatory networks to prioritize future outcomes (Cho et al., 2015; Kelley et al., 2018). Thus, another pathway worth putting forward is that procrastination behavior could be shaped by increasing future task-outcome value. Also, given the integrative nature of prefrontal regulatory functions, we hypothesize a third pathway whereby both decreased task aversiveness and increased task-outcome value may jointly contribute to reduced procrastination, potentially reflecting coordinated downstream effects on valuation and affective processing.

Earlier version: Supporting this view, the left dorsolateral prefrontal cortex (DLPFC), responsible for self-control, has been frequently shown to be closely associated with procrastination. As the biomarker of self-control, the neuroanatomical changes of the left DLPFC were revealed to predict procrastination robustly (Chen et al., 2020; Hu et al., 2018; Liu & Feng, 2017). In addition to structural brain hallmarks, the neurofunctional anomalies underlying procrastination were found in DLPFC-involved circuits (Y. Wu et al., 2016; Xu et al., 2021). Furthermore, the triple brain network theory provides network-based insights to highlight the neural signature of procrastination as the self-control neural network (e.g., left DLPFC; anterior cingulate cortex, ACC), emotional regulation network (e.g., insula and orbitofrontal cortex, OFC) and episodic prospection network (e.g., hippocampus and ventromedial prefrontal cortex, vmPFC, amygdala) (Chen & Feng, 2022; Chen et al., 2020; Schlüter et al., 2018; Wypych et al., 2019). In addition to the left lateralization, there is solid evidence indicating significant associations between self-control and the right DLPFC indeed, particularly given that this region specifically functions in top-down regulation, future self-continuity representation and social decisions (Huang et al., 2025; Knoch & Fehr, 2007; Lin & Feng, 2024). Despite this case, Xu and colleagues demonstrated null effects of anodally stimulating the right DLPFC to modulate either value evaluation or emotional regulation for changing procrastination willingness (Xu et al., 2023). Moreover, a substantial amount of neural evidence supports this conclusion that DLPFC is involved in long-term reward evaluation and value encoding via top-down self-control circuits (Frost & McNaughton, 2017; Jimura et al., 2013; Smith et al., 2018). Using a neurocomputational model, Le Bouc & Pessiglione (2022) provided clear evidence indicating that dorsal PFC signaling expected effort values was significantly attenuated in procrastinators compared to healthy controls (Le Bouc & Pessiglione, 2022). In this vein, this evidence supports this conceptualization that the left DLPFC may be a domain-specific neural signature determining one’s procrastination. In light of technical advances, high-resolution transcranial direct current stimulation (HD-tDCS) has been widely used to reveal the causal neurocognitive mechanism of problematic behaviors by modulating cortical excitability, blood-brain barrier permeability and even neuroplasticity (Cirillo et al., 2017; Woods et al., 2016), which is regulated by the NMDA (N-methyl-D-aspartate) system to either bolster LTP (Long-term potentiation) or LTD (Long-term depression) processes (Chrysikou et al., 2022; Shin et al., 2020). For instance, anodic HD-tDCS applied to the left DLPFC was found effective in inhibiting problematic behaviors caused by the lack of self-control (Allenby et al., 2018), showing significantly amplified local neural oscillations (Chrysikou et al., 2022). Beyond regional neuromodulation in a dose-free protocol, cumulative evidence has well-documented that the effects of tDCS for neuromodulation are highly dose-dependent and are involved in network-wise covariance (Sabé et al., 2024; Soleimani et al., 2023; Woodham et al., 2025). Thus, this study aims to provide causal evidence clarifying the brain-behavior association of procrastination and revealing how self-control works to shape one’s procrastination by manipulating left DLPFC activity in a multiple-session (dose) protocol.

Revised version: Consistent with this framework, the left dorsolateral prefrontal cortex (DLPFC)—a region frequently implicated in value-based decision-making and top-down regulation—has been associated with procrastination. Structural and functional variations in the left DLPFC have been correlated with procrastination tendencies (Chen et al., 2020; Hu et al., 2018; Liu & Feng, 2017). In addition to structural brain hallmarks, the neurofunctional anomalies underlying procrastination were found in DLPFC-involved circuits (Y. Wu et al., 2016; Xu et al., 2021). Furthermore, the triple brain network theory provides network-based insights to highlight the neural signature of procrastination as the self-control neural network (e.g., left DLPFC; anterior cingulate cortex, ACC), emotional regulation network (e.g., insula and orbitofrontal cortex, OFC) and episodic prospection network (e.g., hippocampus and ventromedial prefrontal cortex, vmPFC, amygdala) (Chen & Feng, 2022; Chen et al., 2020; Schlüter et al., 2018; Wypych et al., 2019). In addition to the left lateralization, there is solid evidence indicating significant associations between top-down regulatory process and the right DLPFC indeed, particularly given that this region specifically functions in top-down regulation, future self-continuity representation and social decisions (Huang et al., 2025; Knoch & Fehr, 2007; Lin & Feng, 2024). Despite this case, Xu and colleagues demonstrated null effects of anodally stimulating the right DLPFC to modulate either value evaluation or emotional regulation for changing procrastination willingness (Xu et al., 2023). Moreover, a substantial amount of neural evidence supports this conclusion that DLPFC is involved in long-term reward evaluation and value encoding via top-down regulation circuits (Frost & McNaughton, 2017; Jimura et al., 2013; Smith et al., 2018). Using a neurocomputational model, Le Bouc & Pessiglione (2022) provided clear evidence indicating that dorsal PFC signaling expected effort values was significantly attenuated in procrastinators compared to healthy controls (Le Bouc & Pessiglione, 2022). In this vein, this evidence supports this conceptualization that the left DLPFC may be a domain-specific neural signature determining one’s procrastination. In light of technical advances, high-resolution transcranial direct current stimulation (HD-tDCS) has been widely used to reveal the potential neurocognitive mechanism of problematic behaviors by modulating cortical excitability, blood-brain barrier permeability and even neuroplasticity (Cirillo et al., 2017; Woods et al., 2016), which is regulated by the NMDA (N-methyl-D-aspartate) system to either bolster LTP (Long-term potentiation) or LTD (Long-term depression) processes (Chrysikou et al., 2022; Shin et al., 2020). For instance, anodic HD-tDCS applied to the left DLPFC was found effective in inhibiting problematic behaviors caused by the lack of self-regulatory ability (Allenby et al., 2018), showing significantly amplified local neural oscillations (Chrysikou et al., 2022). Beyond regional neuromodulation in a dose-free protocol, cumulative evidence has well-documented that the effects of tDCS for neuromodulation are highly dose-dependent and are involved in network-wise covariance (Sabé et al., 2024; Soleimani et al., 2023; Woodham et al., 2025). Thus, this study aims to clarify the brain-behavior association between DLPFC neuromodulation and procrastination, and to test whether observed changes in task valuation and aversiveness are consistent with theoretical models of top-down regulation.

Earlier version: To clarify the causal cognitive mechanism of self-control on procrastination, we conducted a double-blind, randomized, multiple-session, placebo-controlled design, with a 2 (active HD-tDCS vs. sham control) × 2 (before first neural stimulation vs. after last neural stimulation) full factorial design (see Fig. 1). This HD-tDCS protocol consisted of 7 sessions spaced over 15 days, and each session was implemented every two days. To ensure sound ecological validity, each procrastinator was informed to report one REAL-LIFE task that he/she should complete the following day (e.g., Day 2) after the current neuromodulation session (e.g., Day 1), and was asked to report the ACTUAL performance for completing this task in that day (Day 2) (see Fig. 1). Based on the temporal decision model of procrastination (Zhang, Liu, et al., 2019), we drew on the experience sampling method (ESM) to estimate the real dynamics of task aversiveness and task-outcome value by using a parameter-free model in each session. More importantly, to clarify which pathway best explains the neurocognitive mechanism of procrastination, we built upon the mixed-effects general linear model and Quasi-Bayesian causal mediation model to test whether the changes in task aversiveness and task outcome value caused by HD-tDCS could predict decreased procrastination. Finally, the follow-up investigation for actual procrastination was conducted for 6 months after the experiment to examine the long-term retention of this neural stimulation effect.

Revised version: To empirically evaluate these theoretical pathways, we conducted a double-blind, randomized, multiple-session, placebo-controlled design, with a 2 (active HD-tDCS vs. sham control) × 2 (before first neural stimulation vs. after last neural stimulation) full factorial design (see Fig. 1). This HD-tDCS protocol consisted of 7 sessions spaced over 15 days, and each session was implemented every two days. To ensure sound ecological validity, each procrastinator was informed to report one REAL-LIFE task that he/she should complete the following day (e.g., Day 2) after the current neuromodulation session (e.g., Day 1), and was asked to report the ACTUAL performance for completing this task in that day (Day 2) (see Fig. 1). Based on the temporal decision model of procrastination (Zhang, Liu, et al., 2019), we drew on the experience sampling method (ESM) to estimate the real dynamics of task aversiveness and task-outcome value by using a parameter-free model in each session. More importantly, to clarify which pathway best explains the neurocognitive mechanisms of procrastination, we built upon the mixed-effects general linear model and Quasi-Bayesian mediation model to test whether the changes in task aversiveness and task outcome value modulated by HD-tDCS could predict decreased procrastination. Finally, the follow-up investigation for actual procrastination was conducted for 6 months after the experiment to examine the long-term after-effect of this neural stimulation.

Experimental diagram of this study.

Earlier version: The upper sub-graph illustrates the whole multiple sessions tDCS neuromodulation pipeline including seven sessions (days) and eight task-demanded days. Here, the “flash” icon indicates conducting tDCS neuromodulation (active anodal stimulation for active NM group and sham stimulation for sham NM group). This “+” label means a task-demanded day, where no stimulation is required for participants but all the covariates of interests should be measured by experience sampling methods. The bottom sub-graph reflect specific pipeline in task-demanded days. Participants were required to provide response at five progressive time moments nearing deadline for task aversiveness and outcome value in task-demanded days. In this diagram, the icon of “clock” symbolize ecological momentary assessment for measuring instant task willingness and outcome value. In addition, twice tests for daily emotions (labeled by “+”) were added for participants at 10:00 and 16:00 as covariates of no interests to be adjusted.

Revised version: The upper sub-graph illustrates the whole multiple sessions tDCS neuromodulation pipeline including seven sessions (days) and eight task-demanded days. Here, the “flash” icon indicates conducting tDCS neuromodulation (active anodal stimulation for active NM group and sham stimulation for sham NM group). This “+” label means a task-demanded day, where no stimulation is required for participants but all the covariates of interest should be measured by experience sampling methods. The bottom sub-graph reflects specific pipeline in task-demanded days. Participants were required to provide response at five progressive time moments nearing deadline for task aversiveness and outcome value in task-demanded days. In this diagram, the icon of “clock” symbolizes ecological momentary assessment for measuring instant task willingness and outcome value. In addition, twice tests for daily emotions (labeled by “+”) were added for participants at 10:00 and 16:00 as covariates of no interests to be adjusted.

Materials and Methods

Earlier version: This study fully adhered to CONSORT reporting guidelines and was originally preregistered in the OSF repository (10.17605/OSF.IO/Y3EDT). However, due to the technical constraint related to OSF account service (see SM Methods), this OSF page is no longer accessible. For transparency and best practices of open science, based on the original protocol documentations, a preregistration statement has been reconstructed to clarify a prior hypotheses, sample size determinations, and analysis plans for this study (Table S1).

Revised version: This study fully adhered to CONSORT reporting guidelines and was originally preregistered in the OSF repository (10.17605/OSF.IO/Y3EDT). However, due to the technical constraint related to OSF account service (see SI Methods), this OSF page is no longer accessible. For transparency and best practices of open science, based on the original protocol documentations, a preregistration statement has been reconstructed to clarify a priori hypotheses, sample size determinations, and analysis plans for this study (Table S1).

Participants

Earlier version: Due to the lack of diagnostic criteria for clinical procrastinators, we recruited a large-scale sample (n = 1,682) to obtain a stable benchmark distribution. Thus, the procrastinators were captured once their procrastination scores were higher than 66 on General Procrastination Scale (GPS) (see Fig. 2A-B). Following this criterion, a total of 186 participants were included initially, which was in accordance with empirical evidence (i.e., 10 - 15% prevalence of procrastination) (Harriott & Ferrari, 1996). Subsequently, the semi-structured interview was performed to screen those suffering from problematic procrastination and volunteering for this study, thereby enrolling 53 participants (see SM Methods). Seven participants were eventually excluded from the analyses because they voluntarily dropped out before experimental completion. All the included participants were screened for depression and anxiety symptoms (see Tab. 1).

Revised version: Due to the lack of diagnostic criteria for clinical procrastinators, we recruited a large-scale sample (n = 1,682) to obtain a stable benchmark distribution. Thus, the procrastinators were captured once their procrastination scores were higher than 66 on General Procrastination Scale (GPS) (see Fig. 2A-B). Following this criterion, a total of 186 participants were included initially, which was in accordance with empirical evidence (i.e., 10 – 15% prevalence of procrastination) (Harriott & Ferrari, 1996). Subsequently, the semi-structured interview was performed to screen those suffering from problematic procrastination and volunteering for this study, thereby enrolling 53 participants (see SI Methods). Seven participants were eventually excluded from the analyses because they voluntarily dropped out before experimental completion. All the included participants were screened for depression and anxiety symptoms (see Tab. 1).

Flow diagram of CONSORT (A) and partial details of randomized groups (B-D) and neural locations of electric pole (E).

Earlier version: B plots the distribution of all the participants procrastination scores (GPS = general procrastination scale). C detailed what the full randomized block design is. D shows the comparison between active neuromodulation group and sham control for procrastination scores. E indicates the pipeline to determine the location of electric pole. The 10-20 EEG standard lead is used to locate the left dorsolaterial prefrontal cortex (DLPFC) initially, and the neuronavigation is further utilized to locate the exact location of this targeting region (i.e., left DLPFC).

Revised version: B plots the distribution of all the participants procrastination scores (GPS = general procrastination scale). C detailed what the full randomized block design is. D shows the comparison between active neuromodulation group and sham control for procrastination scores. E indicates the pipeline to determine the location of electric pole. The 10–20 EEG standard lead is used to locate the left dorsolateral prefrontal cortex (DLPFC) initially, and the neuronavigation is further utilized to locate the exact location of this targeting region (i.e., left DLPFC).

Earlier version: A full randomized block design was used to assign participants to both groups (active neuromodulation group, NM; sham-control group, SC) (see Fig. 2C). As the pilot study probing into the effect of single-session tDCS stimulation to change procrastination willingness indicated (t = 2.38, p = .02, 95% CI [0.14, 1.49]; Xu et al., 2023), statistical power was predetermined by G*Power at a relatively medium effect size (1-β err prob = 0.80, f = 0.25), yielding the total sample size at 18 to reach acceptable power (see SM Methods and Figure. S1). All the participants reported no history for HD-tDCS or neuromodulation. No significant differences were found between groups for any demographic characteristics (see Tab. 1). This study and protocol have been fully approved by the Institutional Review Board (IRB) of the School of Psychology, Southwest University (China, IRB200301108).

Revised version: A full randomized block design was used to assign participants to both groups (active neuromodulation group, NM; sham-control group, SC) (see Fig. 2C). As the pilot study probing into the effect of single-session tDCS stimulation to change procrastination willingness indicated (t = 2.38, p = .02, 95% CI [0.14, 1.49]; Xu et al., 2023), statistical power was predetermined by G*Power at a relatively medium effect size (1-β err prob = 0.80, f = 0.25), indicating the total sample size at 34 (17 per group) to reach acceptable power. To account for potential attrition, we determined to recruit 36 participants, at least. (see SI Methods and Fig. S1). All the participants reported no history for HD-tDCS or neuromodulation. No significant differences were found between groups for any demographic characteristics (see Tab. 1). This study and protocol have been fully approved by the Institutional Review Board (IRB) of the School of Psychology, Southwest University (China, IRB200301108).

Measurement

Earlier version: The General Procrastination Scale (GPS) developed by Lay (1986) was used to quantify one’s chronic procrastination symptom here (Lay, 1986). This scale was widely adopted in many cross-cultural contexts, and has been reported to have good psychometric properties (Klein et al., 2019). There were two additional items that we added for lie detection, including description of “the sky is red” and “I have never taken shower”. If either one was selected as “agree” by participant, his/her response would be discarded. Internal reliability of the GPS in the current study was found acceptable (Cronbach’s α = 0.890). No significant difference was found between groups for GPS scores (t = -1.08, 95% CI: -4.296 1.283, p = .283; Jeffreys-Zellner-Siow Bayesian factor, BF, BF10 = 0.455, error = 0.020 %).

Revised version: The General Procrastination Scale (GPS) developed by Lay (1986) was used to quantify one’s chronic procrastination symptom here (Lay, 1986). This scale was widely adopted in many cross-cultural contexts, and has been reported to have good psychometric properties (Klein et al., 2019). There were two additional items that we added for lie detection, including description of “the sky is red” and “I have never taken a shower”. If either one was selected as “agree” by participant, his/her response would be discarded. Internal reliability of the GPS in the current study was found acceptable (Cronbach’s ± = 0.890). No significant difference was found between groups for GPS scores (t = −1.08, 95% CI: −4.296 1.283, p = .283; Jeffreys-Zellner-Siow Bayesian factor, BF, BF10 = 0.455, error = 0.020 %).

Removed in the revised version: To quantify one’s daily emotions, the Positive and Negative Affect Schedule (PANAS) developed by Watson et al. (1988) was adopted. This scale consists of two subscales (i.e., positive affect, PA and negative affect, NA), each containing 10 terms, where higher scores indicate stronger affective states (Terracciano et al., 2003). To control for potential confounding effects of daily emotions, both subscale scores were included as covariates of no interest in the statistical model. Additionally, this tool demonstrated good internal consistency for both sub-scales (α = 0.801 for PA, α = 0.793 for NA).

Experimental design and procedure

Nested cross-sectional longitudinal design

Earlier version: This study used a nested cross-sectional longitudinal design to investigate whether the multiple-session anodal HD-tDCS targeting the left DLPFC could reduce actual procrastination behavior and to probe how this effect manifests. To assess procrastination in daily life, we implemented a 15-day protocol alternating between Neuromodulation Days (Days 2, 4, 6, 8, 10, 12, 14) and Task Days (Days 1, 3, 5, 7, 9, 11, 13, 15). On the Neuromodulation days, the 20-min anodal HD-tDCS neuromodulation targeting the left DLPFC was performed for HD-tDCS active group at intervals of 2 days, while the sham-control group received sham HD-tDCS training. This HD-tDCS training was repeated for a total of seven sessions, and lasted 15 days (see Fig. 1). Crucially, to capture procrastination in ecologically valid contexts, prior to receiving either active or sham HD-tDCS (administered between 09:00-18:00), participants were instructed to specify a real-life task they were personally obligated to complete the following day, with a self-defined deadline strictly constrained to 18:00-24:00 to ensure ≥24 hours between stimulation offset and task deadline, thereby isolating offline after-effects. This task should meet the following three criteria: (a) it should be already assigned in the real-world settings; (b) deadline should be constrained to 18:00-24:00 (see above); (c) it should be more likely to induce procrastinate. By doing so, more than 300 real-life tasks were collected, spanning academic (e.g., “submit a statistics homework assignment”), occupational (e.g., “draft and email a project proposal”), administrative (e.g., “complete online application for Class C driver’s license”), self-improvement (e.g., “practice guitar for ≥ 30 minutes”), domestic (e.g., “do laundry”), and health-related (e.g., “running 2,000m for exercise”). Full task list has been tabulated in the Appendix 1. As primary outcomes, all the participants were required to report task-execution willingness (TEW) (Zhang & Feng, 2020; Zhang, Liu, et al., 2019), for a real-life task 24 hours post-neuromodulation. Thus, procrastination willingness was quantified as 100-TEW score (see underneath for details). Furthermore, we asked participants to report the actual task completion rate (CR) of the task at the deadline (e.g. participant A finished 90% homework at deadline and reported this situation to us at deadline). In this vein, the actual procrastination rate (PR) was quantified as 1-CR.

Revised version: This study used a nested cross-sectional longitudinal design to investigate whether the multiple-session anodal HD-tDCS targeting the left DLPFC could reduce actual procrastination behavior and to probe how this effect manifests. To assess procrastination in daily life, we implemented a 15-day protocol alternating between Neuromodulation Days (Days 2, 4, 6, 8, 10, 12, 14) and Task Days (Days 1, 3, 5, 7, 9, 11, 13, 15). On the Neuromodulation days, the 20-min anodal HD-tDCS neuromodulation targeting the left DLPFC was performed for the HD-tDCS active group at intervals of 2 days, while the shamcontrol group received sham HD-tDCS training. This HD-tDCS training was repeated for a total of seven sessions, and lasted 15 days (see Fig. 1). Crucially, to capture procrastination in ecologically valid contexts, priori to receiving either active or sham HD-tDCS (administered between 09:00-18:00), participants were instructed to specify a real-life task they were personally obligated to complete the following day, with a self-defined deadline strictly constrained to 18:00-24:00 to ensure ≥24 hours between stimulation offset and task deadline, thereby isolating offline after-effects. This task should meet the following three criteria: (a) it should be already assigned in the real-world settings; (b) deadline should be constrained to 18:00-24:00 (see above); (c) it should be more likely to induce procrastinate. By doing so, more than 300 real-life tasks were collected, spanning academic (e.g., “submit a statistics homework assignment”), occupational (e.g., “draft and email a project proposal”), administrative (e.g., “complete online application for Class C driver’s license”), selfimprovement (e.g., “practice guitar for ≥30 minutes”), domestic (e.g., “do laundry”), and health-related (e.g., “running 2,000m for exercise”). Full task list has been tabulated in the Appendix 1. As primary outcomes, all the participants were required to report task-execution willingness (TEW) (Zhang & Feng, 2020; Zhang, Liu, et al., 2019), for a real-life task 24 hours post-neuromodulation. Thus, procrastination willingness was quantified as 100-TEW score (see underneath for details). Furthermore, we asked participants to report the actual task completion rate (CR) of the task at the deadline (e.g. participant A finished 90% homework at deadline and reported this situation to us at deadline). In this vein, the actual procrastination rate (PR) was quantified as 1-CR.

On the Task Day, we developed a mobile app to implement experience sampling method (ESM) for tracking one’s real-time evaluation of task aversiveness and task outcome value (see Fig. 1). The task aversiveness describes how disagreeable one perceives performing a given real-life task to be, whereas outcome value refers to the subjective benefits of the task outcome brought about by completing the task before the deadline (Zhang & Feng, 2020). As theoretically conceptualized by the temporal decision model (TDM) of procrastination, the perceived task aversiveness is hyperbolically discounted when approaching deadline, showing sharply discounting when faring away from deadline but slowly discounting once nearing deadline (Zhang & Feng, 2020). Thus, considering this nonlinear dynamics inherent in this hyperbolic discounting, the five recording moments of ESM were selected per task a priori by using a log-spaced temporal sampling scheme (Myerson et al., 2001), with increasing sampling density toward the deadline, such as moments of 10:00 (earliest), 16:00, 18:00, 19:30, 20:00 (deadline). The five sampling points could meet statistical prerequisite in the hyperbolic model fitting, requiring ≥ 4 points (Green & Myerson, 2004). To do so, recording moments of tasks were individually tailored for each task per participant in this ESM procedure. To obviate the confounds of daily emotions in task aversiveness evaluation, we used the averaged scores of PANAS at 10:00 (noon) and 16:00 (afternoon) as anchoring points to quantify one’s daily emotions by using this ESM app. Before each session of HD-tDCS training, each participant was required to report a real-life task whose deadline is tomorrow. To obtain the long-term effect of HD-tDCS (i.e., the interval between HD-tDCS and task completion is at least 24 hours), the task deadline that participants reported was required to be between 18:00 – 24:00. Once a sampling time reached, this app would send a digital message to require participants to fill online form for data collection.

Quantification of covariates of interests

Earlier version: Outcome variables of this study were twofold: one is task-execution willingness and another is procrastination rate (PR). Task-execution willingness is used to evaluate one’s subjective inclination to avoid procrastination (Zhang & Feng, 2020). In this vein, we used a 100-point scale to require participants to report their task-execution willingness (0 for “I will definitely procrastinate this task” and 100 for “I will take action to complete this task immediately”). This metric was recorded 24 hours after neuromodulation to examine its long-term effects. PR is used to quantify the extent to which one task has been procrastinated, and was calculated as 1 - CR (task completion rate). Critically, at the precise deadline, the app prompted participants to (a) indicate task completion status (yes/no), and if incomplete, (b) report the percentage completed (1-99%), defined as the Task CR, while simultaneously uploading objective evidence (e.g., screenshots of submitted files, photos of physical outputs, system-generated logs, or app-exported records). If the task was actually completed before the deadline, the CR would be 100% and the PR would be calculated as 0% (1-CR). PR was recorded at the actual task deadline for each participant. We were also interested in re-investigating their actual procrastination by using PR 6 months after the last neuromodulation to test the long-term retention of this neuromodulation effect.

Revised version: Outcome variables of this study were twofold: one is task-execution willingness and another is procrastination rate (PR). Task-execution willingness is used to evaluate one’s subjective inclination to avoid procrastination (Zhang & Feng, 2020). In this vein, we used a 100-point scale to require participants to report their task-execution willingness (0 for “I will definitely procrastinate this task” and 100 for “I will take action to complete this task immediately”). This metric was recorded 24 hours after neuromodulation to examine its long-term effects. PR is used to quantify the extent to which one task has been procrastinated, and was calculated as 1 − CR (task completion rate). Critically, at the precise deadline, the app prompted participants to (a) indicate task completion status (yes/no), and if incomplete, (b) report the percentage completed (1-99%), defined as the Task CR, while simultaneously uploading objective evidence (e.g., screenshots of submitted files, photos of physical outputs, system-generated logs, or app-exported records). If the task was actually completed before the deadline, the CR would be 100% and the PR would be calculated as 0% (1-CR). PR was recorded at the actual task deadline for each participant. We were also interested in re-investigating their actual procrastination by using PR 6 months after the last neuromodulation to test the long-term after-effect of this neuromodulation.

Earlier version: From what has been mentioned above, task aversiveness and outcome value were considered key factors to explain the effect of neuromodulation on reducing procrastination in the current study. To quantify one’s task aversiveness, participants were required to rate their feelings towards the task by using a 100-point visual analog scale (i.e., How do you feel in the current moment when you need to complete this task before deadline, with 0 for “extremely unpleasant”, 50 for “totally neutral” and 100 for “extremely pleasant”). Likewise, participants are also required to rate their outcome value by using a 100-point visual analog scale (i.e., How much do you desire to obtain incentive outcome of this task, with 0 for “extremely weak”, 50 for “medium” and 100 for “extremely strong”). As articulated temporal decision theoretical model above, the task aversiveness evoked by executing a task was temporally dynamic in a hyperbolic discounting pattern, with sharply discounting in faring away from deadline but slowly discounting in nearing deadline (Zhang & Feng, 2020). To quantitatively characterize the task aversiveness with consideration for its dynamics, the model-free area under the curve (AUC) was calculated. Specifically, based on the log-spaced temporal sampling rule, task aversiveness was measured by 100-point visual analog scale at the five sampling moments. Then, the task aversiveness discounting (A) was calculated as 1- (A(t) / A(earliest)), where t(earliest) was the earliest sampling point, serving as the reference for immediate execution. Subsequently, using the GraphPad Prism software (v9, 525), the AUC was computed as the trapezoidal integration between task aversiveness discounting and time across five data points, basing on the Myerson algorithm (Myerson et al., 2001). By doing so, a higher AUC reflects stronger temporal discounting of task aversiveness along with nearing deadline, which means that participants experience a faster decline in subjective aversiveness as execution is delayed, yielding lower effective aversiveness and reduced avoidance behavior. As for the task outcome value, it was theoretically posited as a relatively stable evaluation of the task (Zhang & Feng, 2020). Therefore, it was quantified by the self-reported 100-point visual analog scale after neuromodulation at least 12 hours later, to ensure no online effect.

Revised version: From what has been mentioned above, task aversiveness and outcome value were considered key factors to explain the effect of neuromodulation on reducing procrastination in the current study. To quantify one’s task aversiveness, participants were required to rate their feelings towards the task by using a 100-point visual analog scale (i.e., How do you feel in the current moment when you need to complete this task before deadline, with 0 for “extremely unpleasant”, 50 for “totally neutral” and 100 for “extremely pleasant”). Likewise, participants are also required to rate their outcome value by using a 100-point visual analog scale (i.e., How much do you desire to obtain incentive outcome of this task, with 0 for “extremely weak”, 50 for “medium” and 100 for “extremely strong”). As articulated temporal decision theoretical model above, the task aversiveness evoked by executing a task was temporally dynamic in a hyperbolic discounting pattern, with sharply discounting in faring away from deadline but slowly discounting in nearing deadline (Zhang & Feng, 2020). To quantitatively characterize the task aversiveness with consideration for its dynamics, the model-free area under the curve (AUC) was calculated. Specifically, based on the log-spaced temporal sampling rule, task aversiveness was measured by 100-point visual analog scale at the five sampling moments. Then, the task aversiveness discounting (A) was calculated as 1- (A(t) / A(earliest)), where t(earliest) was the earliest sampling point, serving as the reference for immediate execution. Subsequently, using the GraphPad Prisma software (v9, 525), the AUC was computed as the trapezoidal integration between task aversiveness discounting and time across five data points, basing on the Myerson algorithm (Myerson et al., 2001). By doing so, a higher AUC reflects stronger temporal discounting of task aversiveness along with nearing deadline, which means that participants experience a faster decline in subjective aversiveness as execution is delayed, yielding lower effective aversiveness and reduced avoidance behavior. As for the task outcome value, it was theoretically posited as a relatively stable evaluation of the task (Zhang & Feng, 2020). Therefore, it was quantified by the selfreported 100-point visual analog scale after neuromodulation at least 12 hours later, to ensure no online effect.

HD-tDCS protocol

The HD-tDCS suit (stimulater and 4 × 1 multichannel stimulation adapter, MSA) that this study used was produced by the Soterix Medical Inc., and has been widely verified safe, effective and reliable for public (Villamar et al., 2013). Based on advanced properties of 4 × 1 MSA, the targeted areas for current flow can be constrained within 2.5 cm2 (Villamar et al., 2013).

Earlier version: To position electrodes into targeted areas (left DLPFC), the 10-20 international system for EEG was initially used to mark potential nodes, and determined Cz as reference point. There is compelling evidence to claim that the left F3 could be used as target node for modulating the left DLPFC (Seibt et al., 2015; Tsukuda et al., 2025). In this vein, the central anodal electrode was determined onto F3, and four return electrodes surrounded central electrode at outside of 7.5 cm, including F5, AF3, FC3 and F1 (see Fig. 2F). Ramp-up and ramp-down durations were set to 30 seconds. To further locate the targeted areas, the high-definition neuronavigation system (ANT Neuro Inc., Welbergweg, Germany) was performed. Results indicated the accurate position that we pre-determined by showing a high overlap probability over the left DLPFC (MNI Coordinate: -51 40 18, 94.33 % overlapping probability) (see SM Method and Fig. 2F). In addition, the coordinates of this targeted area were retrieved from the Brede Database (http://neuro.imm.dtu.dk/services/), and showed highly pertinent functions related to DLPFC.

Revised version: To position electrodes into targeted areas (left DLPFC), the 10-20 international system for EEG was initially used to mark potential nodes, and determined Cz as reference point. There is compelling evidence to claim that the left F3 could be used as target node for modulating the left DLPFC (Seibt et al., 2015; Tsukuda et al., 2025). In this vein, the central anodal electrode was determined onto F3, and four return electrodes surrounded central electrode at outside of 7.5 cm, including F5, AF3, FC3 and F1 (see Fig. 2E). Ramp-up and ramp-down durations were set to 30 seconds. To further locate the targeted areas, the high-definition neuronavigation system (ANT Neuro Inc., Welbergweg, Germany) was performed. Results indicated the accurate position that we pre-determined by showing a high overlap probability over the left DLPFC (MNI Coordinate: −51 40 18, 94.33 % overlapping probability) (see SI Method and Fig. 2E). In addition, the coordinates of this targeted area were retrieved from the Brede Database (http://neuro.imm.dtu.dk/services/), and showed highly pertinent functions related to DLPFC.

Before stimulation, participants were informed to clean the scalp to reduce resistance. Then, cotton swabs were used to separate hair until the scalp surface become visible. Subsequently, the electrically conductive gel (about 1.5 ml) was introduced into the plastic casing facilitating constraint of current flow. Next, the Ag/AgCI sintered ring electrodes were placed onto plastic casing and covered with a cap to lock them in the right positions. To reduce discomfort, the electrode cables were taped elsewhere. Further, the above processes would be re-adjusted if any electrode resistances were found larger than 1.5 units (Villamar et al., 2013). Once all the processes had been completed, the stimulator would be launched.

Participants in the HD-tDCS training group underwent constant electric current of 2.0 mA targeting the left DLPFC for 20 minutes. Results from the simulation of electric density showed a peak current of ~0.5mA/cm2 at the central electrode and of ~0.125 mA/cm2 at the four return electrodes, thereby indicating the safety and effectiveness. As for the sham-controlled group, the stimulator would deliver current flow with 2.0 mA during the first and last 30 seconds to elicit sense of electric stimulation for blinding of them. To obtain the pure offline effect, these measures for task-execution willingness, task aversiveness, and outcome value were conducted after stimulation at least 12 hours (Bikson et al., 2016).

Statistics

All the statistics were implemented by R (https://www.rstudio.com/) and R-dependent packages.

Earlier version: To clarify whether multiple-session HD-tDCS neuromodulation can reduce procrastination, the generalized mixed-effects linear model (GLMM) was constructed with full factorial design for subjective procrastination willingness (i.e., self-reported visual analog scores) and actual procrastination behavior (i.e., real-world task-completion rate before deadline). Here, sex, age and socioeconomic status (SES) were modeled as covariates of no interest. As the National Bureau of Statistics (China) issued (https://www.stats.gov.cn/sj/tjbz/gjtjbz/), on the basis of per capita annual household income, the SES was divided into seven hierarchical tiers from 1 (poor) to 7 (rich). To obviate subjective rating bias stemming from individual daily mood, we separately measured participants’ daily emotional fluctuation at 10:00 and 16:00 using a self-rating visual analog item (i.e., “How do feel for your mood today?”, 0 for “completely uncomfortable” and 100 for “definitely happy”). By doing so, the averaged score of those self-rating emotions at the two time points was modeled into the GLMM as covariate of no interest, yielding the final expression of “outcome ~ Group*Treatment_Day + Age + Gender + SES + Emotions + (1 + Treatment_Day | SubjectID)” in the statistical model”. This analysis was implemented using the “lme4” and “lmerTest” packages. Employing “emmeans” package, simple effects were also tested at baseline and post-last-intervention using Tukey-adjusted pairwise comparisons of estimated marginal means from the full GLMM, controlling for covariates and random-effects structure. To validate statistical robustness, instead of continuous outcomes for parametric tests, we also conducted a between-group comparison for the number of tasks that procrastination emerges by using the nonparametric x2 test with φ correction or Fisher exact test. Regarding the 6-month follow-up investigation, this GLMM was also built to examine the long-term retention of neuromodulation on reducing actual procrastination.

Revised version: To clarify whether multiple-session HD-tDCS neuromodulation can reduce procrastination, the general linear mixed-effects model (LMM) was constructed for subjective procrastination willingness (i.e., self-reported visual analog scores) and actual procrastination behavior (i.e., real-world task-completion rate before deadline). Here, sex, age and socioeconomic status (SES) were modeled as covariates of no interest. As the National Bureau of Statistics (China) issued (https://www.stats.gov.cn/sj/tjbz/gjtjbz/), on the basis of per capita annual household income, the SES was divided into seven hierarchical tiers from 1 (poor) to 7 (rich). To obviate subjective rating bias stemming from individual daily mood, we separately measured participants’ daily emotional fluctuation at 10:00 and 16:00 using a self-rating visual analog item (i.e., “How do you feel today?”, 0 for “completely uncomfortable” and 100 for “definitely happy”). By doing so, the averaged score of those self-rating emotions at the two time points was modeled into the LMM as a covariate of no interest. Given the risks of convergence failure with the two correlated random-effects structure (i.e., treatment days and self-reported emotions), we hypothesized that the random effects are independent, leading to model simplification. Consistent with this assumption, model comparisons favored the simplified independent structure over the full correlated structure for both outcomes. For the actual procrastination model, the simplified model showed lower AIC (Δ = −1.0) and BIC (Δ = −11.8), with a non-significant likelihood ratio test (χ2(3) = 5.04, p = .17). For the task-execution willingness model, the simplified model was also preferred (ΔAIC = −5.5, ABIC = −16.4; LRT: χ2(3) = 0.51, p = .91). This analysis was implemented using the “lme4” and “lmerTest” packages. Employing “emmeans” package, simple effects were also tested at baseline and post-last-intervention using Tukey-adjusted pairwise comparisons of estimated marginal means from the full LMM, controlling for covariates and random-effects structure. To validate statistical robustness, instead of continuous outcomes for parametric tests, we also conducted a between-group comparison for the number of tasks that procrastination emerges by using the nonparametric x2 test with φ correction or Fisher exact test. Regarding the 6- month follow-up investigation, this LMM was also built to examine the long-term after-effect of neuromodulation on reducing actual procrastination. To examine whether our findings were sensitive to the distributional assumptions of the dependent variables, we re-analyzed the main models using an alternative distributional specification. Given that both procrastination rates (ranging from 0% to 100%) and task-execution willingness (measured on a 0–100 visual analog scale) are bounded continuous outcomes, and that procrastination rates frequently reached the lower boundary (0%) in the present study, a Beta regression model with a logit link function was employed for a sensitivity analysis.

Earlier version: To ascertain the neurocognitive mechanism of tDCS in reducing procrastination, the Quasi-Bayesian causal mediation analysis was used to model the association between the effects of tDCS, task aversiveness/outcome and decreased procrastination. To build upon this model, the tDCS treatments were inputted as independent variables, and the task aversiveness/outcomes were modeled as causal mediating variables by using the “Mediation” package (https://cran.r-project.org/web/packages/mediation/) (Imai et al., 2010). We estimated these pathway effects (i.e., averaged causal mediation effects, δ; averaged direct effects, ζ; total effects, ρ) by using Markov Chain Monte Carlo (MCMC) sampling. To improve the statistical reliability, the sequential ignorability assumption was tested by using sensitivity analysis. Details for the statistical principals and basis could be found elsewhere (Imai et al., 2010).

Revised version: To ascertain the neurocognitive mechanism of tDCS in reducing procrastination, the Quasi-Bayesian mediation analysis was used to model the association between the effects of tDCS, task aversiveness/outcome and decreased procrastination. To build upon this model, the tDCS treatments were inputted as independent variables, and the task aversiveness/outcomes were modeled as mediating variables by using the “Mediation” package (https://cran.r-project.org/web/packages/mediation/) (Imai et al., 2010). We estimated these pathway effects (i.e., average mediation effects, δ; average direct effect, ζ; total effects, ρ) by using Markov Chain Monte Carlo (MCMC) sampling. To improve the statistical reliability, the sequential ignorability assumption was tested by using sensitivity analysis. Details for the statistical principals and basis could be found elsewhere (Imai et al., 2010). As these mediation analyses are based on observational measures rather than experimentally manipulated mediators, they do not establish causal pathways. Simulation studies have shown that such analyses can suffer from inflated Type 1 error rates. Results should therefore be interpreted as hypothesis-generating and statistically consistent with the proposed theoretical model, rather than as confirmatory evidence of causal mechanisms.

Results

Blinding

In both groups, almost all participants reported perceiving acceptable pain stemming from current stimulation, and believed they were receiving treatment, with 91.30% (21/23) for active neuromodulation group (NM) and with 86.95% (20/23) for sham control group (SC) (X2 = 0.224, p = .636). All the participants were engaged in the identical experimental procedures excepting stimulation’s type (active vs sham). In addition, statistical models excluded Session 1 and Session 4 because participants reported additional unexpected events that uncontrollably disrupt task execution in both groups (see SI Result and Tab. S2).

Multiple-sessions HD-tDCS (ms-tDCS) can alleviate procrastination

Earlier version: To identify whether ms-tDCS targeting the left DLPFC can alleviate subjective procrastination willingness and actual procrastination behavior, a generalized linear mixed-effects model (GLMM) with Satterwhite algorithm was built, with task-execution willingness and actual procrastination rates (PR) as primary outcomes, respectively. For procrastination willingness, results showed a statistically significant interaction effect between multi-session neuromodulations and groups (β = -7.8, SE = 1.8, DF = 45.6, p < .001; Fig. 3A). In the post-hoc simple effect analysis, it demonstrated a significantly increased task-execution willingness (i.e., decreased procrastination willingness) after neuromodulation in the active neuromodulation group (NM-before: 35.65 ± 30.20, NM-after: 80.43 ± 19.92, t.ratio = 5.4, p < .0001, Tukey correction), but no such effects were identified in the sham control group (SC-before: 37.57 ± 26.46, SC-after: 47.35 ± 30.49, t.ratio = 0.3, p = .77, Tukey correction) (Fig. 3B-C). A linear uptrend for task-execution willingness was further observed across multiple sessions in the active NM group, indicating gradually increasing neuromodulation effects (Fig. 3D; p < .01, Mann-Kendall test). For actual procrastination behavior, changes to actual procrastination rates across all the sessions have been detailed in the Fig. 3E. Similarly, a statistically significant interaction effect was identified here (β = -7.4, SE = 2.4, DF = 46.6, p = .004), and the simple effect analysis further revealed decreased actual procrastination rates after ms-tDCS in the active neuromodulation group (NM-before: 43.26 ± 39.09, NM-after: 0.00 ± 0.00, t.ratio = 5.1, p < .0001, Tukey correction), but no such prominent changes found in the sham control group (SC-before: 46.47 ± 40.75, SC-after: 33.34 ± 37.82, t.ratio = 0.7, p = .48, Tukey correction) (Fig. 3F-G). Also, a significant downtrend for procrastination rates across all the sessions was identified in the active NM group (Fig. 3H; p < .01, Mann-Kendall test).

Revised version: To identify whether ms-tDCS targeting the left DLPFC can alleviate subjective procrastination willingness and actual procrastination behavior, a general mixed-effects linear model (LMM) with Satterthwaite’s method was built, with task-execution willingness and actual procrastination rates (PR) as primary outcomes, respectively. For procrastination willingness, results showed a statistically significant interaction effect between multi-session neuromodulations and groups (β = −7.84, SE = 1.80, t = −4.36, DF = 45.6, p < .001; Fig. 3A). In the post-hoc simple effect analysis, it demonstrated a significantly increased task-execution willingness (i.e., decreased procrastination willingness) after neuromodulation in the active neuromodulation group (NM-before: 35.65 ± 30.21, NM-after: 80.43 ± 19.92, Mean Diff = 41.79, SE = 7.58, DF = 103.4, t.ratio = 5.51, p < .0001, Tukey correction), but no such effects were identified in the sham control group (SC-before: 37.57 ± 26.46, SC-after: 47.35 ± 30.49, Mean Diff = 2.58, SE = 7.56, DF = 96.8, t.ratio = 0.34, p = .73, Tukey correction) (Fig. 3B-C). A linear uptrend for task-execution willingness was further observed across multiple sessions in the active NM group, indicating gradually increasing neuromodulation effects (Fig. 3D; p < .01, Mann-Kendall test). For actual procrastination behavior, changes to actual procrastination rates across all the sessions have been detailed in the Fig. 3E. Similarly, a statistically significant interaction effect was identified here (β = −7.37, SE = 2.40, t = −3.02, DF = 46.6, p = .004), and the simple effect analysis further revealed decreased actual procrastination rates after ms-tDCS in the active neuromodulation group (NM-before: 56.74 ± 39.10, NM-after: 0.00 ± 0.00, Mean Diff = 44.40, SE = 9.36, DF = 110.0, t.ratio = 4.74, p < .0001, Tukey correction), but no such prominent changes found in the sham control group (SC-before: 46.47 ± 40.76, SC-after: 33.35 ± 37.82, Mean Diff = 7.53, SE = 9.28, DF = 102.0, t.ratio = 0.81, p = .42, Tukey correction) (Fig. 3F-G). Also, a significant downtrend for procrastination rates across all the sessions was identified in the active NM group (Fig. 3H; p < .01, Mann-Kendall test).

Results of neuromodulation effects to task-execution willingness and procrastination rates (PR).

Earlier version: A shows the effects of neuromodulation to increase task-execution willingness for both active group and sham control across sessions that included in formal analysis (session 0 (baseline), 2, 3, 5, 6, 7). B illustrates the effects of whole neuromodulation round to task-execution willingness for both group. C plots the changes of task-execution willingness for both group after neuromodulation. D provides a line chart to show the changes of task-execution willingness across each session that included in formal analysis. E shows the effects of neuromodulation to reduce PR for both active group and sham control across sessions that included in formal analysis. F illustrates the effects of whole neuromodulation round to PR for both group. G plots the absolute changes of PR for both group after neuromodulation. H provides a line chart to show the changes of PR across these sessions that included in formal analysis. Pie graph for each comparison represent the corresponding result of Bayesian factor inference, with brown piece for supporting H1 evidence and white piece for supporting H0 evidence. Each bar indicates mean value, and each line placed onto the bar reflects standard deviation (SD).

Revised version: A shows the effects of neuromodulation to increase task-execution willingness for both active group and sham control across sessions that included in formal analysis (session 0 (baseline), 2, 3, 5, 6, 7). B illustrates the effects of whole neuromodulation round to task-execution willingness for both groups. C plots the changes of task-execution willingness for both groups after neuromodulation. D provides a line chart to show the changes of task-execution willingness across each session that included in formal analysis. E shows the effects of neuromodulation to reduce PR for both active group and sham control across sessions that included in formal analysis. F illustrates the effects of whole neuromodulation round to PR for both groups. G plots the absolute changes of PR for both groups after neuromodulation. H provides a line chart to show the changes of PR across these sessions that included in formal analysis. Pie graph for each comparison represents the corresponding result of Bayesian factor inference, with brown piece for supporting H1 evidence and white piece for supporting H0 evidence. Each bar indicates mean value, and each line placed onto the bar reflects standard deviation (SD).

Earlier version: Furthermore, the nonparametric x2 test of R × C contingency table was conducted for the count of procrastinated tasks, by treating this outcome (i.e., whether a participant actually procrastinates the task in real-world settings) as an ordinal variable. Results showed a significant reduction in group-averaged procrastination frequency in the active neuromodulation group after last-session HD-tDCS, but not in the sham control group (NM-before: 69.56% (16/23 participants), NM-after: 0.00% (0/23 participants); SC-before: 69.56% (16/23 participants), SC-after: 56.52% (13/23 participants), x2= 10.08, p < .001), partly indicating a high statistical robustness. To systematically test the statistical robustness, we reconstructed GLMMs by iteratively removing data from the last two sessions, which showed extraordinarily high effectiveness from neuromodulation. Results showed the significant group*neuromodulation_sessions interaction effects across all those nested models (removing session #6, #7 or both, all p < .05; see SI Results and Tab. S3-4). In brief, these findings provided empirical evidence to support that ms-tDCS neuromodulation targeting the left DLPFC can be an effective way to reduce both procrastination willingness and actual procrastination.

Revised version: Furthermore, the nonparametric X2 test of R × C contingency table was conducted for the count of procrastinated tasks, by treating this outcome (i.e., whether a participant actually procrastinates the task in real-world settings) as an ordinal variable. Results showed a significant reduction in group-averaged procrastination frequency in the active neuromodulation group after last-session HD-tDCS, but not in the sham control group (NM-before: 69.56% (16/23 participants), NM-after: 0.00% (0/23 participants); SC-before: 69.56% (16/23 participants), SC-after: 56.52% (13/23 participants), X2= 10.08, p< .001), partly indicating a high statistical robustness. To systematically test whether such effects are biased by extreme data points or patterns, we reran the main LMMs by iteratively removing data from the last two sessions, which showed extraordinarily high effectiveness from neuromodulation (e.g., all the participants in the neuromodulation group had no actual procrastination behavior in session #6 and #7). Results showed the significant group*neuromodulation sessions interaction effects across all those nested models (removing session #6, #7 or both, all p < .05; see SI Results and Tab. S3–4), potentially indicating a statistical robustness from the data pattern. Furthermore, as a sensitivity analysis, we reran the main LMMs using Beta regression distribution with a logit link function, which is appropriate for both bounded outcomes mentioned above (i.e., procrastination rate and procrastination willingness). The main effects (i.e., Group and Treatment day) and their interaction remained significant for both procrastination rate and willingness (see SI Results and Tab. S5), confirming that our findings are robust to alternative distributional assumptions. In brief, these findings provided empirical evidence to support that ms-tDCS neuromodulation targeting the left DLPFC can be an effective way to reduce both procrastination willingness and actual procrastination.

Ms-tDCS changes task aversiveness and task-outcome value

Earlier version: Both task aversiveness and task outcome value serve as key pathways determining whether one would procrastinate. To this end, we further utilized a generalized linear mixed-effects model to examine the effects of ms-tDCS on changes in task aversiveness and task outcome value. Task aversiveness changes across all the sessions are shown in the Fig. 4A and 4C. We demonstrated a statistically significant decrease in task aversiveness and an increase in task outcome value via ms-tDCS in the neuromodulation group (Task aversiveness: interaction effect, β = -0.12, SE = 0.04, DF = 46.7, p = .002; simple effect, NM-before(AUC): 1.13 ± 0.53, NM-after(AUC): 1.95 ± 0.85, t.ratio = 4.5, p < .001, Tukey correction; Outcome value: β = -6.8, SE = 1.74, DF = 46.2, p< .001; simple effect, NM-before: 35.86 ± 27.82, NM-after: 73.08 ± 23.33, t.ratio = 5.0, p < .001, Tukey correction; see Fig. 4B), but not in the sham control group (Task aversiveness: SC-before (AUC): 1.07 ± 0.51, SC-after (AUC): 1.28 ± 0.46, t.ratio = 1.3, p = .20, Tukey correction; Outcome value: SC-before: 34.00 ± 25.17, SC-after: 40.13 ± 28.94, t.ratio = 0.8, p = .41, Tukey correction; see Fig. 4D). In the neuromodulation (NM) group, task aversiveness steadily decreased with the cumulative number of stimulation sessions, while perceived task outcome value increased significantly (see Fig. 4E-F, p < .05, Mann-Kendall test). Thus, it provides causal evidence clarifying that neuromodulation to left DLPFC reduces task aversiveness and enhances task-outcome value meanwhile.

Revised version: Both task aversiveness and task outcome value serve as key pathways determining whether one would procrastinate. To this end, we further utilized a generalized linear mixed-effects model to examine the effects of ms-tDCS on changes in task aversiveness and task outcome value. Task aversiveness changes across all the sessions are shown in the Fig. 4A and 4C. We demonstrated a statistically significant decrease in task aversiveness and an increase in task outcome value via ms-tDCS in the neuromodulation group (Task aversiveness: interaction effect, β = −0.12, SE = 0.04, DF = 46.69, t = −3.28, p = .002; simple effect, NM-before(AUC): 1.13 ± 0.54, NM-after(AUC): 1.95 ± 0.84, Mean Diff = 0.81, SE = 0.16, DF = 104.5, t.ratio = 4.91, p < .001, Tukey correction; Outcome value: β = −6.76, SE = 1.74, DF = 46.2, t = −3.89, p< .001; simple effect, NM-before: 35.87 ± 27.83, NM-after: 73.09 ± 23.34, Mean Diff = 39.36, SE = 7.25, DF = 103.7, t.ratio = 5.43, p < .001, Tukey correction; see Fig. 4B), but not in the sham control group (Task aversiveness: SC-before(AUC): 1.07 ± 0.51, SC-after(AUC): 1.28 ± 0.46, Mean Diff = 0.17, SE = 0.17, DF = 98.3, t.ratio = 1.06, p= .29, Tukey correction; Outcome value: SC-before: 34.00 ± 25.17, SC-after: 40.13 ± 28.94, Mean Diff = 5.54, SE = 7.26, DF = 98.7, t.ratio = 0.76, p = .44, Tukey correction; see Fig. 4D). In the neuromodulation (NM) group, task aversiveness steadily decreased with the cumulative number of stimulation sessions, while perceived task outcome value increased significantly (see Fig. 4E-F, p < .05, Mann-Kendall test). Thus, these findings are consistent with the view that neuromodulation of the left DLPFC is associated with reduced task aversiveness and increased task-outcome value.

Results of neuromodulation effects to task aversiveness and outcome value.

Earlier version: A shows the effects of neuromodulation to increase AUC of task aversiveness for both active group and sham control across sessions that included in formal analysis (session 0 (baseline), 2, 3, 5, 6, 7). Higher AUC indicates lower task avsiveness as a given task is increasingly executed. B plots the changes of AUC of task aversiveness for both group after neuromodulation. C shows the effects of neuromodulation to increase outcome value for both active group and sham control across sessions that included in formal analysis. D plots the changes of outcome value for both group after neuromodulation. E provides a line chart to show the changes of AUC of task aversiveness across these sessions that included in formal analysis. F provides a line chart to show the changes of outcome value in the same manner. Pie graph for each comparison represent the corresponding result of Bayesian factor inference, with brown piece for supporting H1 evidence and white piece for supporting H0 evidence. Each bar indicates mean value, and each line placed onto the bar reflects standard deviation (SD).

Revised version: A shows the effects of neuromodulation to increase AUC of task aversiveness for both active group and sham control across sessions that included in formal analysis (session 0 (baseline), 2, 3, 5, 6, 7). Higher AUC indicates lower task aversiveness as a given task is increasingly executed. B plots the changes of AUC of task aversiveness for both groups after neuromodulation. C shows the effects of neuromodulation to increase outcome value for both active group and sham control across sessions that included in formal analysis. D plots the changes of outcome value for both groups after neuromodulation. E provides a line chart to show the changes of AUC of task aversiveness across these sessions that included in formal analysis. F provides a line chart to show the changes of outcome value in the same manner. Pie graph for each comparison represents the corresponding result of Bayesian factor inference, with brown piece for supporting H1 evidence and white piece for supporting H0 evidence. Each bar indicates mean value, and each line placed onto the bar reflects standard deviation (SD).

Increased task outcome value but not decreased task aversiveness predicts reduced procrastination

Earlier version: Given the dual neurocognitive pathways identified above—reduced task aversiveness and increased task-outcome value—we proposed that these changes, conceptually driven by enhanced self-control via ms-tDCS over left DLPFC, account for how neuromodulation reduces procrastination. To this end, we utilized a generalized linear model to regress decreased task aversiveness and increased task outcome value to changes in task-execution willingness. In this model, increased task outcome value (Δ Outcome value) significantly predicted increased task-execution willingness (Δ task-execution willingness) (x2 = 15.95, p <.01, R2 = .40; Δ Outcome value:β= 0.61, S.E. = 0.12, p< .001, 95% CI: 0.37 - 0.86), whereas no significant effect was observed for predicting task-execution willingness through decreased task aversiveness (Δ task aversiveness:β= 0.10, S.E. = 0.12, p= .41, 95% CI: -0.14 - 0.34). Likewise, for actual procrastination behavior in real-world settings, increased outcome value (Δ Outcome value) was identified to be significantly predictive, whereas decreased task aversiveness showed null effects (see Tab. 2). Collectively, these findings provide causal evidence supporting notion that the outcome-value neurocognitive pathway accounts well for procrastination reduction, rather than a dual circuits incorporating task-aversiveness.

Revised version: Given the dual neurocognitive pathways identified above—reduced task aversiveness and increased task-outcome value—we proposed that these changes may reflect downstream consequences of prefrontal neuromodulation on value-based decision processes, statistically consistent with theoretical models of top-down regulation. To this end, we utilized a general linear model to regress decreased task aversiveness and increased task outcome value to changes in task-execution willingness. In this model, increased task outcome value (ΔOutcome value) significantly predicted increased task-execution willingness (Δtask-execution willingness) (X2 = 15.95, p <.01, R2 = .40; ΔOutcome value: β = 0.61, S.E. = 0.12, p < .001,95% CI: 0.37 – 0.86), whereas no significant effect was observed for predicting task-execution willingness through decreased task aversiveness (ΔTask aversiveness:β= 0.10, S.E. = 0.12, p = .41, 95% CI: −0.14 – 0.34). Likewise, for actual procrastination behavior in real-world settings, increased outcome value (ΔOutcome value) was identified to be significantly predictive, whereas decreased task aversiveness showed null effects (see Tab. 2). Collectively, these findings provide statistical evidence consistent with the hypothesis that the outcome-value pathway may contribute to procrastination reduction.

Added in the revised version: Increased task outcome value is specifically associated with reduced procrastination in the contexts of neuromodulation

Removed in the revised version: Increased task outcome value plays causal role to explain why ms-tDCS reduces procrastination

Earlier version: To clarify the causal neurocognitive mechanism of procrastination, the Quasi-Bayesian causal mediation analysis was undertaken by using White’s heteroskedasticity-consistent estimator, with increased task outcome value as a mediated variable. Results demonstrated the significant causal mediated role of increased task outcome value in increasing the task-execution willingness (δ = 21.73, p < .01; ζ = 11.25, p = .07, ρ = 32.99, p < .01, simulation = 1,000; see Fig. 5A) and real-world procrastination (δ = 30.75, p < .01; ζ = 3.05, p = .52, ρ = 33.81, p < .01, simulation = 1,000; see Fig. 5B), as caused by ms-tDCS neuromodulation. To ensure the robustness and specificity of these findings, the sensitivity analysis was implemented by changing sampling parameters and outcome variables. By doing so, those findings were validated highly robust, as shown by replicated observations across bootstrapping sampling algorithms (see SI Results and Tab. S5-6). Moreover, the results of the control analysis further validated the specificity of these findings by showing a null causal mediated effect of this model to predict one’s task aversiveness (see SI Results and Tab. S7). In summary, these findings demonstrated a mechanistic pathway underlying procrastination: the self-control that was conceptualized to be governed by left DLPFC mitigate procrastination by plausibly increasing task-outcome value.

Revised version: To explore the potential neurocognitive pathways of procrastination, the Quasi-Bayesian mediation analysis was undertaken, with increased task outcome value as a statistically mediated variable. As an exploratory analysis, results indicated that increased task outcome value was associated with changes in the task-execution willingness (δ = 21.73, p < .01; ζ = 11.25, p = .07, ρ = 32.99, p < .01, simulation = 1,000; see Fig. 5A) and real-world procrastination (δ = 30.75, p < .01; ζ = 3.05, p = .52, ρ = 33.81, p <.01, simulation = 1,000; see Fig. 5B), in the context of ms-tDCS neuromodulation. To ensure the statistical robustness and specificity of these findings, the sensitivity analysis was implemented by changing sampling parameters and outcome variables. By doing so, those findings were validated statistically robust, as shown by replicated observations across bootstrapping sampling subsets (see SI Results and Tab. S6–7). Moreover, the results of the control analysis further validated the specificity of these findings by showing a null statistically mediated effect of this model to predict one’s task aversiveness (see SI Results and Tab. S8). In summary, these findings identified a statistical pathway consistent with the theoretical model: neuromodulation of the left DLPFC was associated with increased task-outcome value, which in turn was associated with reduced procrastination. Nevertheless, all mediation findings are now explicitly labeled as “exploratory” and framed as quantifying statistical associations consistent with the TDM pathway, not causal mediation.

Earlier version: Results of Quasi-Bayesian causal model for the medicated role of increased task outcome value in the association between neuromodulation and task-execution willingness (A) and actual procrastination rates (B).

Revised version: Results of Quasi-Bayesian model for the mediated role of increased task outcome value in the association between neuromodulation and task-execution willingness (A) and actual procrastination rates (B).

Earlier version: ADE = averaged direct effect, ACME = averaged causal mediated effect, CI = confident interval.

Revised version: ADE = average direct effect, ACME = average causal mediation effect, CI = confidence interval.

Long-term effects of ms-tDCS

Earlier version: We have also attempted to conduct a follow-up investigation to test the long-term retention of ms-tDCS in reducing actual procrastination. Almost all the participants had undergone follow-up except one in the neuromodulation group after last neuromodulation for 6 months (NNM = 22, NSC = 23). Thus, the GLMM was constructed, with the PR before first neuromodulation vs. PR after last neuromodulation for 6 months as covariates of interest. Results showed the statistically significant group*time interaction effects (β = 16.5, SE = 9.9, p = .049). Simple-effect model demonstrated a decrease in actual procrastination rates in the active neuromodulation group after last stimulation for 6 months compared to baseline (β = -22.05, SE = 10.0, p = .038, Tukey correction; NM-before: 40.68 ± 37.96, NM-after6-months: 18.63 ± 29.80), and revealed null effects in the SC group (β= 1.26, SE = 9.78, p = .99, Tukey correction; SC-before: 46.47 ± 40.75, SC-after6-months: 47.73 ± 39.18) (see Fig. 6).. Furthermore, using a nonparametric x2 test to compare differences in the number of procrastinated tasks, we still found a statistically significant reduction in procrastination frequency in NM group after neuromodulation for 6 months compared to baseline (x2 = 3.30, p = .035, NM-before: 68.19% (15/22), NM-after6-months: 40.91% (9/22)), while no significant changes were observed in the SC group (x2 = 0.11, p = .74, SC-before: 69.56% (16/23), SC-after6-months: 73.91% (17/23)). Therefore, beyond short-term effects, the benefits of ms-tDCS neuromodulation to reduce procrastination pose the long-term retention.

Revised version: We have also attempted to conduct a follow-up investigation to test the long-term after-effect of ms-tDCS in reducing actual procrastination. Almost all the participants had undergone follow-up except one in the neuromodulation group after last neuromodulation for 6 months (NNM = 22, NSC = 23). Thus, the LMM was constructed, with the PR before first neuromodulation vs. PR after last neuromodulation for 6 months as covariates of interest. Results showed the statistically significant group*time interaction effects (β = 16.5, SE = 9.9, p = .049). Simple-effect model demonstrated a decrease in actual procrastination rates in the active neuromodulation group after last stimulation for 6 months compared to baseline (β = - 22.05, SE = 10.0, p= .038, Tukey correction; NM-before: 40.68 ± 37.96, NM-after6-months: 18.63 ± 29.80), and revealed null effects in the SC group (β= 1.26, SE = 9.78, p= .99, Tukey correction; SC-before: 46.47 ± 40.75, SC-after6-months: 47.73 ± 39.18) (see Fig. 6). Furthermore, using a nonparametric X2 test to compare differences in the number of procrastinated tasks, we still found a statistically significant reduction in procrastination frequency in NM group after neuromodulation for 6 months compared to baseline (X2 = 3.30, p = .035, NM-before: 68.19% (15/22), NM-after6-months: 40.91% (9/22)), while no significant changes were observed in the SC group (X2 = 0.11, p = .74, SC-before: 69.56% (16/23), SC-after6-months: 73.91% (17/23)). Therefore, beyond short-term effects, the benefits of ms-tDCS neuromodulation on reducing procrastination were still detectable at a 6-month follow-up, providing preliminary evidence consistent with long-term after-effects.

Changes of actual procrastination rates among pre-test, post-test and 6-month follow-up for active neuromodulation group and sham-control group.

Earlier version: Pre-test means to test the actual procrastination rates before first HD-tDCS neuromodulation. Post-test means to test the actual procrastination rates after last neuromodulation. The 6-month follow means to re-investigate the actual procrastination rates after last neuromodulation for 6 months.

Revised version: Pre-test means to test the actual procrastination rates before first HD-tDCS neuromodulation. Post-test means to test the actual procrastination rates after last neuromodulation. The 6-month follow-up means to re-investigate the actual procrastination rates after last neuromodulation for 6 months.

Discussion

Earlier version: In the current study, by performing anodal ms-tDCS neuromodulation on the left DLPFC, both procrastination willingness and actual procrastination behavior were significantly decreased in real life. Additionally, a 6-month follow-up investigation revealed the long-term retention of such effects. Furthermore, this neuromodulation was found to decrease task aversiveness and increase outcome values; notably, only increased task-outcome value could predict decreased procrastination. On balance, our findings clarified the neurocognitive mechanism of procrastination by showing that self-control could increase task-outcome value so as to reduce procrastination. In addition, this study provided an effective way to reduce actual procrastination by using ms-tDCS neuromodulation.

Revised version: In the current study, by performing anodal ms-tDCS neuromodulation on the left DLPFC, both procrastination willingness and actual procrastination behavior were significantly decreased in real life. Additionally, a 6-month follow-up investigation revealed the long-term retention of such effects. Furthermore, this neuromodulation was found to decrease task aversiveness and increase outcome values; notably, only increased task-outcome value could predict decreased procrastination. On balance, our findings provide evidence consistent with the hypothesis that neuromodulation of the left DLPFC is associated with reduced procrastination, primarily through increasing task-outcome value rather than merely reducing task aversiveness. In addition, this study provided an effective way to reduce actual procrastination by using ms-tDCS neuromodulation.

Earlier version: One major contribution this study has made is to disentangle the neurocognitive mechanism of procrastination by demonstrating that self-control could increase task-outcome value so as to reduce procrastination. Neurobiological substrates of procrastination have been investigated in recent years, and have demonstrated the crucial roles of the left DLPFC in predicting procrastination (Chen & Feng, 2022; Chen et al., 2021; Chen et al., 2022; Hu et al., 2018; Liu & Feng, 2017; Zhang et al., 2017; Zhang et al., 2016). Not only the brain functional anomalies of the left DLPFC but also the neuroanatomical disruptions of self-control brain network constituted by the left DLPFC were found to be linked with more procrastination behaviors (Chen & Feng, 2022; Zhang et al., 2016). Notwithstanding this, it still remains unclear to claim their causal brain-behavior relationship - that is - no known evidence existed to clarity whether changes of the left DLPFC lead to procrastination or vice versa. The current study demonstrated the causal role of the left DLPFC in procrastination by showing that the neuromodulation of the left DLPFC indeed manipulated procrastination, and thus provided straightforward and powerful evidence to fill this gap.

Revised version: One contribution of this study is to provide empirical evidence partially consistent with the temporal decision model (TDM), showing that increased task-outcome value—rather than decreased task aversiveness—was statistically associated with reduced procrastination following DLPFC neuromodulation. Neurobiological substrates of procrastination have been investigated in recent years, and have demonstrated the crucial roles of the left DLPFC in predicting procrastination (Chen & Feng, 2022; Chen et al., 2021; Chen et al., 2022; Hu et al., 2018; Liu & Feng, 2017; Zhang et al., 2017; Zhang et al., 2016). Not only the brain functional anomalies of the left DLPFC but also the neuroanatomical disruptions of self-regulatory brain network constituted by the left DLPFC were found to be linked with more procrastination behaviors (Chen & Feng, 2022; Zhang et al., 2016). Notwithstanding this, it still remains unclear to claim their brain-behavior relationship - that is - no known evidence existed to clarity whether changes of the left DLPFC lead to procrastination or vice versa. The current study demonstrated the role of the left DLPFC in procrastination by showing that the neuromodulation of the left DLPFC indeed manipulated procrastination, and thus provided straightforward and powerful evidence to fill this gap.

Earlier version: It has long been acknowledged that the left DLPFC is associated with many aspects of self-control communities, such as patience to wait long-term gratification for a delay, inhibition of impulsiveness, and control of game addiction (Cohen & Lieberman, 2010; Lin & Feng, 2024). Furthermore, the increased activation of the left DLPFC has been observed during the exertion of self-control which modulates value signals (Hare et al., 2009; Harris et al., 2013). Meanwhile, procrastination has been argued to be the consequence of self-control failures and self-regulation failure for a long time (Ariely & Wertenbroch, 2002; Rebetez et al., 2018; Rozental & Carlbring, 2014). Supporting this, both the brain morphological disruptions in the DLPFC and anomalies in the functional coupling of the DLPFC-based self-control network were interpreted as the phenotype of reduced self-control, making individuals more prone to procrastinate tasks, as reported in existing literature (Xu et al., 2021; Yang et al., 2021). Moreover, it was worth noting that the increased activation of the left DLPFC was found to be involved in outcome value evaluation through self-control regulation (Chen et al., 2018; Zha et al., 2019). There was more straightforward evidence to substantiate the role of manipulating the DLPFC in changing one’s subjective value evaluation (Huang et al., 2017; Minati et al., 2012). On the other hand, the theoretical explanations and empirical evidence have increasingly converged into one line for claiming that the increased task outcome value would prompt more motivation to drive one to take action immediately and thus reduce procrastination (Zhang, Liu, et al., 2019). In this vein, this study advanced our understanding of the neurocognitive mechanism of procrastination by showing that the cortical excitability of the DLPFC produced by active neuromodulation could boost self-control to increase task outcome value so as to reduce procrastination behavior.

Revised version: It has long been acknowledged that the left DLPFC is consistently implicated in top-down regulatory processes and value-based decision-making, such as patience to wait long-term gratification for a delay, inhibition of impulsiveness, and control of game addiction (Cohen & Lieberman, 2010; Lin & Feng, 2024). Furthermore, the increased activation of the left DLPFC has been observed during the exertion of self-regulation which modulates value signals (Hare et al., 2009; Harris et al., 2013). Meanwhile, procrastination has been argued to be the consequence of self-regulation failure for a long time (Ariely & Wertenbroch, 2002; Rebetez et al., 2018; Rozental & Carlbring, 2014). Supporting this, both the brain morphological disruptions in the DLPFC and anomalies in the functional coupling of the DLPFC-based regulatory network have been correlated with procrastination tendencies (Xu et al., 2021; Yang et al., 2021). Moreover, it was worth noting that the increased activation of the left DLPFC was found to be involved in outcome value evaluation through self-control regulation (Chen et al., 2018; Zha et al., 2019). There is more straightforward evidence to substantiate the role of manipulating the DLPFC in changing one’s subjective value evaluation (Huang et al., 2017; Minati et al., 2012). Moreover, the theoretical explanations and empirical evidence have increasingly converged into one line for claiming that the increased task outcome value would prompt more motivation to drive one to take action immediately and thus reduce procrastination (Zhang et al., 2019). Building on this foundation, among several theoretical interpretations and cognitive pathways, our study showed the one plausible neurocognitive mechanism of procrastination: the cortical excitability of the DLPFC produced by active neuromodulation may engage prefrontal regulatory networks to increase task outcome value, which in turn is associated with reduced procrastination behavior, statistically supporting the theoretical accounts of temporal decision model (TDM, Zhang et al., 2019). TDM uniquely posits procrastination as contingent on the trade-off between task aversiveness and taskoutcome value; our observation that only increased outcome value—not decreased aversiveness—statistically predicted reduced procrastination aligns precisely with TDM’s hypothesis that value-based processes may dominate affective-avoidance processes in driving behavioral change. However, neither TMT (which emphasizes temporal discounting of utility per se) nor the mood repair perspective (which prioritizes short-term affect regulation) explicitly predicts this dissociable pattern. Despite statistically supporting the TDM, we acknowledge that alternative neurocognitive mechanisms could contribute to the observed reductions in procrastination. For instance, repeated exposure to the experience-sampling protocol may have enhanced participants’ awareness of task progress or facilitated feedbackbased learning, thereby increasing the subjective value of goal completion independent of DLPFC neuromodulation. Similarly, improvements in working memory for task maintenance, attentional allocation to reported goals, or strategic shifts in goal selection across sessions could plausibly mediate the intervention effects. While our double-blind, sham-controlled design and inclusion of daily emotional covariates help mitigate some non-specific confounds, the present study did not incorporate direct measures of these alternative processes. Consequently, we cannot definitively isolate the value-based pathway posited by the TDM from concurrent contributions of attention, learning, or executive functions.

Earlier version: In addition, another contribution of the current study is to provide an effective way to reduce both procrastination willingness and actual procrastination in real-life tasks. As mentioned above, despite the fact that multifarious behavioral interventions and evidence have been massively studied for overcoming procrastination, they have shared a common aim - that is - reducing the intention-action gap (Miao et al., 2024; van Hooft et al., 2005). Eerde and Klingsieck (2018) put forward an insightful standpoint established by meta-analytic evidence: procrastination is characterized as an intention-action gap rather than an intention to postpone (van Eerde & Klingsieck, 2018). This notion has been partly supported by showing that cognitive behavior therapy (CBT) for goal-directed behaviors may outperform other interventions focusing on time management (Rozental et al., 2018). Moreover, the trans-theoretical model of procrastination has shown that the behavioral intervention may be effective in changing one’s motivation to overcome procrastination but not in actual behaviors (Grunschel & Schopenhauer, 2015). Thus, the ms-tDCS neuromodulation the current study performed has a remarkable advantage in reducing procrastination - that is, the intention-action gap was attenuated so as to overcome procrastination behavior effectively. Furthermore, both 2-day-interval long-term effects and the 6-month long-term retention of the effects of ms-tDCS on reducing actual procrastination have been revealed as well. Thus far, the trends in adopting tDCS neuromodulation techniques in many aspects of behavioral therapies have emerged, but concern for a long-lasting effect of single session stimulation has continued (Brunoni et al., 2013; Brunoni et al., 2012). To tackle this concern well, instead of single-session tDCS, the current study adopted multiple-session stimulation to implement neuromodulation on the left DLPFC, which facilitates long-term effects (Au et al., 2017; Tedesco Triccas et al., 2016). Existing neurobiological theories and empirical evidence have demonstrated that multiple-session tDCS stimulation could boost cumulative effects of consolidation for activity-dependent LTP (long-term potentiation), which is crucial to neurobehavioral learning, and thus produce robust long-term after-effects (Agboada et al., 2020; Au et al., 2017). Intriguingly, the activity-dependent LTP process produced by multiple-session consolidation was found to contribute to long-term cortical plasticity, especially in the DLPFC (Jannati et al., 2023; Siebner & Rothwell, 2003). Thus, it led us to presume that the 6-month long-term retention of the ms-tDCS effect on reducing procrastination might be attributed to long-term neuroplastic changes in the DLPFC. On balance, this study provided an effective way to help procrastinators overcome actual procrastination in real-life.

Revised version: In addition, another contribution of the current study is to provide an effective way to reduce both procrastination willingness and actual procrastination in real-life tasks. As mentioned above, despite the fact that multifarious behavioral interventions and evidence have been massively studied for overcoming procrastination, they have shared a common aim - that is - reducing the intention-action gap (Miao et al., 2024; van Hooft et al., 2005). Eerde and Klingsieck (2018) put forward an insightful standpoint established by meta-analytic evidence: procrastination is characterized as an intention-action gap rather than an intention to postpone (van Eerde & Klingsieck, 2018). This notion has been partly supported by showing that cognitive behavior therapy (CBT) for goal-directed behaviors may outperform other interventions focusing on time management (Rozental et al., 2018). Moreover, the trans-theoretical model of procrastination has shown that the behavioral intervention may be effective in changing one’s motivation to overcome procrastination but not in actual behaviors (Grunschel & Schopenhauer, 2015). Thus, the ms-tDCS protocol employed here may offer a promising avenue for reducing procrastination by attenuating the intention-action gap, though further validation in larger, clinically screened cohorts is warranted. Furthermore, both 2-day-interval long-term effects and the 6-month long-term retention of the effects of ms-tDCS on reducing actual procrastination have been revealed as well. Thus far, the trends in adopting tDCS neuromodulation techniques in many aspects of behavioral therapies have emerged, but concern for a long-lasting effect of single session stimulation has continued (Brunoni et al., 2013; Brunoni et al., 2012). To tackle this concern well, instead of single-session tDCS, the current study adopted multiple-session stimulation to implement neuromodulation on the left DLPFC, which facilitates long-term effects (Au et al., 2017; Tedesco Triccas et al., 2016). Existing neurobiological theories and empirical evidence have demonstrated that multiplesession tDCS stimulation could boost cumulative effects of consolidation for activitydependent LTP (long-term potentiation), which is crucial to neurobehavioral learning, and thus produce robust long-term after-effects (Agboada et al., 2020; Au et al., 2017). Intriguingly, the activity-dependent LTP process produced by multiple-session consolidation was found to contribute to long-term cortical plasticity, especially in the DLPFC (Jannati et al., 2023; Siebner & Rothwell, 2003). Thus, the detectable effects at 6 months are consistent with the hypothesis that repeated neuromodulation may induce neuroplastic changes in the DLPFC that support sustained behavioral change. However, we explicitly note that a single follow-up timepoint cannot establish the stability or trajectory of these effects; future studies with multiple longitudinal assessments are required to substantiate claims about long-term retention. On balance, this study provided an effective way to help procrastinators overcome actual procrastination in real-life.

Limitations

Earlier version: While the use of a multi-session design and the long-term assessment can be considered a strength of the present study, it also has several limitations. Even though various tDCS effects have been demonstrated so far, they also tend to be difficult to replicate and sensitive to not yet fully understood context conditions and interindividual differences, which also applies to transcranial magnetic stimulation (TMS) (Valle et al., 2009). To overcome this shortcoming, it will be necessary to establish an individually tailored tDCS protocol to improve the sensitivity of corresponding interventions (Chew et al., 2015). Thus, future research could further improve the effects of tDCS on reducing procrastination by adopting more individualized tDCS protocols. Another limitation is the lack of real-time functional neuroimaging measures to better monitor the impact of our intervention. In the absence of such measures, we had to rely on behavioral indicators to assess the success of the tDCS training. In addition to technical limitations, given the apparently limited size of the sample (total N = 46), it warrants caution in generalizing these findings elsewhere, and necessitates further validations in a large-scale cohort. Also, considering the lack of medical screening for psychiatric conditions (e.g., ADHD or depression) in this sample, it remains unclear whether these training effects are domain-specific for procrastination. Moreover, this study did not collect data for assessing participants’ self-control at either baseline or post-neuromodulation, thereby limiting our ability to determine whether the effects on procrastination were uniquely attributable to neuromodulation-induced changes in self-control. In addition, despite instructing to report valid real-life tasks with high probabilities to procrastinate, we had not yet measured the task difficulty and consistency across sessions for each participant. Consequently, interpreting the effects of neuromodulation to mitigate procrastination as “unique contributions” should warrant cautions. Finally, given the absence of cathodal HD-tDCS stimulation as a contrast condition in causal inference, it warrants caution that the increased DLPFC excitability may not be the exclusively neural mechanism for procrastination.

Revised version: While the use of a multi-session design and the long-term assessment can be considered a strength of the present study, it also has several limitations. Even though various tDCS effects have been demonstrated so far, they also tend to be difficult to replicate and sensitive to not yet fully understood context conditions and interindividual differences, which also applies to transcranial magnetic stimulation (TMS) (Valle et al., 2009). To overcome this shortcoming, it will be necessary to establish an individually tailored tDCS protocol to improve the sensitivity of corresponding interventions (Chew et al., 2015). Thus, future research could further improve the effects of tDCS on reducing procrastination by adopting more individualized tDCS protocols. Another limitation is the lack of real-time functional neuroimaging measures to better monitor the impact of our intervention. In the absence of such measures, we had to rely on behavioral indicators to assess the success of the tDCS training. In addition to technical limitations, given the apparently limited size of the sample (total N = 46), it warrants caution in generalizing these findings elsewhere, and necessitates further validations in a large-scale cohort. Also, considering the lack of medical screening for psychiatric conditions (e.g., ADHD or depression) in this sample, it remains unclear whether these training effects are domainspecific for procrastination. Notably, we explicitly acknowledge that exploratory Quasi-Bayesian mediation analyses, based on observational measures rather than experimentally manipulated mediators, cannot establish causal pathways and are susceptible to inflated Type 1 error rates as demonstrated in recent simulation studies. These findings should be interpreted strictly as hypothesis-generating and statistically consistent with the proposed theoretical model, rather than as confirmatory evidence of causal mechanisms. Substantiating the precise neurocognitive pathway will require future studies employing stronger causal designs, such as experimental manipulation of task valuation or longitudinal cross-lagged modeling. Moreover, this study did not collect data for assessing participants’ self-control at either baseline or post-neuromodulation. Accordingly, we explicitly note that self-control was not directly measured or manipulated in this study; the observed associations between DLPFC neuromodulation, task-outcome value, and procrastination are consistent with theoretical models positing a role for top-down regulatory processes, but do not constitute direct evidence that self-control mechanisms were engaged. This limitation precludes definitive conclusions about the unique contribution of self-control-related pathways versus alternative neurocognitive mechanisms. In addition, despite instructing to report valid real-life tasks with high probabilities to procrastinate, we did not measure the task difficulty and consistency across sessions for each participant. Consequently, interpreting the effects of neuromodulation to mitigate procrastination as “unique contributions” warrants caution. Given the absence of cathodal HD-tDCS stimulation as a contrast condition in inference, it also warrants caution that the increased DLPFC excitability may not be the exclusive neural mechanism for procrastination. Finally, while our 6-month follow-up provides preliminary evidence for sustained effects, the use of a single follow-up timepoint limits our ability to characterize the temporal trajectory of retention. Future studies incorporating multiple followup assessments (e.g., 1-month, 3-month, 6-month, 12-month) would strengthen evidence for the robustness and durability of this intervention.

Conclusions

Earlier version: In conclusion, this study potentially provides an effective way to reduce both procrastination willingness and actual procrastination behavior by using neuromodulation on the left DLPFC. Furthermore, such effects have been observed for 2-day-interval long-term after-effects, and were also found for 6-month long-term retention in part. More importantly, this study identified that the ms-tDCS neuromodulation could decrease task aversiveness and increase task outcome value, and further demonstrated that the increased task outcome value could predict decreased procrastination, a relationship conceptually driven by enhancing self-control. In this vein, the current study enriches our understanding of neurocognitive mechanism of procrastination by showing the prominent role of increased task outcome value in reducing procrastination. Also, it may provide an effective method for intervening in human procrastination.

Revised version: In conclusion, this study potentially provides an effective way to reduce both procrastination willingness and actual procrastination behavior by using neuromodulation on the left DLPFC. Furthermore, such effects have been observed for 2-day-interval long-term after-effects, and were also found for 6-month long-term retention in part. More importantly, this study identified that the ms-tDCS neuromodulation could decrease task aversiveness and increase task outcome value, and further demonstrated that the increased task outcome value could predict decreased procrastination, a relationship that is conceptually consistent with theoretical models of top-down regulation and value-based decision-making. In this vein, the current study enriches our understanding of neurocognitive mechanism of procrastination by showing the prominent role of increased task outcome value in reducing procrastination. Also, it may inform the development of theory-driven, neuromodulation-informed strategies for behavioral interventions, pending further validation in diverse populations.